Flower/FLWR-1 regulates neuronal activity via the plasma membrane Ca2+-ATPase to promote recycling of synaptic vesicles

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

The authors investigated the role of the C. elegans Flower protein, FLWR-1, in synaptic transmission, vesicle recycling, and neuronal excitability. They confirmed that FLWR-1 localizes to synaptic vesicles and the plasma membrane and facilitates synaptic vesicle recycling at neuromuscular junctions. They observed that hyperstimulation results in endosome accumulation in flwr-1 mutant synapses, suggesting that FLWR-1 facilitates the breakdown of endocytic endosomes. Using tissue-specific rescue experiments, the authors showed that expressing FLWR-1 in GABAergic neurons restored the aldicarb-resistant phenotype of flwr-1 mutants to wild-type levels. By contrast, cholinergic neuron expression did not rescue aldicarb sensitivity at all. They also showed that FLWR-1 removal leads to increased Ca2+ signaling in motor neurons upon photo-stimulation. From these findings, the authors conclude that FLWR-1 helps maintain the balance between excitation and inhibition (E/I) by preferentially regulating GABAergic neuronal excitability in a cell-autonomous manner.

Overall, the work presents solid data and interesting findings, however the proposed cell-autonomous model of GABAergic FLWR-1 function may be overly simplified in my opinion.

Most of my previous comments have been addressed; however, two issues remain.

(1) I appreciate the authors' efforts conducting additional aldicarb sensitivity assays that combine muscle-specific rescue with either cholinergic or GABergic neuron-specific expression of FLWR-1. In the revised manuscript, they conclude, "This did not show any additive effects to the pure neuronal rescues, thus FLWR-1 effects on muscle cell responses to cholinergic agonists must be cellautonomous." However, I find this interpretation confusing for the reasons outlined below.

Figure 1 - Figure Supplement 3B shows that muscle-specific FLWR-1 expression in flwr-1 mutants significantly restores aldicarb sensitivity. However, when FLWR-1 is co-expressed in both cholinergic neurons and muscle, the worms behave like flwr-1 mutants and no rescue is observed. Similarly, cholinergic FLWR-1 alone fails to restore aldicarb sensitivity (shown in the previous manuscript).

This data is still shown in the manuscript, Fig. 3D. We interpreted our finding in the muscle/cholinergic co-rescue experiment as meaning, that FLWR-1 in cholinergic neurons over-compensates, so worms should be resistant, and the rescuing effect of muscle FLWR-1 is therefore cancelled. But it is true, if this were the case, why does the pure cholinergic rescue not show over-compensation? We added a sentence to acknowledge this inconsistency and we added a sentence in the discussion (see also below, comment 1) of reviewer #2).

These observations indicate a non-cell-autonomous interaction between cholinergic neurons and muscle, rather than a strictly muscle cell-autonomous mechanism. In other words, FLWR-1 expressed in cholinergic neurons appears to negate or block the rescue effect of muscle-expressed FLWR-1. Therefore, FLWR-1 could play a more complex role in coordinating physiology across different tissues. This complexity may affect interpretations of Ca2+ dynamics and/or functional data, particularly in relation to E/I balance, and thus warrants careful discussion or further investigation.

For the Ca2+ dynamics, we think the effects of flwr-1 are likely very immediate, as the imaging assay relies on a sensor expressed directly in the neurons or muscles under study, and not on indirect phenotypes as muscle contraction and behavior, that depend on an interplay of several cell types influencing each other.

(2) The revised manuscript includes new GCaMP analyses restricted to synaptic puncta. The authors mention that "we compared Ca2+ signals in synaptic puncta versus axon shafts, and did not find any differences," concluding that "FLWR-1's impact is local, in synaptic boutons." This is puzzling: the similarity of Ca2+ signals in synaptic regions and axon shafts seems to indicate a more global effect on Ca2+ dynamics or may simply reflect limited temporal resolution in distinguishing local from global signals due to rapid Ca2+ diffusion. The authors should clarify how they reached the conclusion that FLWR-1 has a localized impact at synaptic boutons, given that synaptic and axonal signals appear similar. Based on the presented data, the evidence supporting a local effect of FLWR-1 on Ca2+ dynamics appears limited.

We apologize, here we simply overlooked this misleading wording in our rebuttal letter. The data we mentioned, showing no obvious difference in axon vs. bouton, are shown below, including time constants for the onset and the offset of the stimulus (data is peak normalized for better visualization):

One can see that axonal Ca2+ signals may rise a bit slower than synaptic Ca2+ signals, as expected for Ca2+ entering the boutons, and then diffusing out into the axon. The loss of FLWR1 does not affect this. However, the temporal resolution of the used GCaMP6f sensor is ca. 200 ms to reach peak, and the decay time (to t1/2) is ca. 400 ms (PMID: 23868258). Thus, it would be difficult to see effects based on Ca2+ diffusion using this assay. For the decay, this is similar for both axon and synapse, while flwr-1 mutants do not reduce Ca2+ as much as wt. In the axon, there is a seemingly slightly slower reduction in flwr-1 mutants, however, given the kinetics of the sensor, this is likely not a meaningful difference. Therefore, we wrote we did not find differences. The interpretation should not have been that the impact of FLWR-1 is local. It may be true if one could image this at faster time scales, i.e. if there is more FLWR-1 localized in boutons (as indicated by our data showing FLWR-1 enrichment in boutons; Fig. 3), and when considering its possible effect on MCA-3 localization (and assuming that MCA-3 is the active player in Ca2+ removal), i.e. FLWR-1 recruiting MCA-3 to boutons (Fig. 9C, D).

Reviewer #2 (Public review):

Summary:

The Flower protein is expressed in various cell types, including neurons. Previous studies in flies have proposed that Flower plays a role in neuronal endocytosis by functioning as a Ca2+ channel. However, its precise physiological roles and molecular mechanisms in neurons remain largely unclear. This study employs C. elegans as a model to explore the function and mechanism of FLWR-1, the C. elegans homolog of Flower. This study offers intriguing observations that could potentially challenge or expand our current understanding of the Flower protein. Nevertheless, further clarification or additional experiments are required to substantiate the study's conclusions.

Strengths:

A range of approaches was employed, including the use of a flwr-1 knockout strain, assessment of cholinergic synaptic activity via analyzing aldicarb (a cholinesterase inhibitor) sensitivity, imaging Ca2+ dynamics with GCaMP3, analyzing pHluorin fluorescence, examination of presynaptic ultrastructure by EM, and recording postsynaptic currents at the neuromuscular junction. The findings include notable observations on the effects of flwr-1 knockout, such as increased Ca2+ levels in motor neurons, changes in endosome numbers in motor neurons, altered aldicarb sensitivity, and potential involvement of a Ca2+-ATPase and PIP2 binding in FLWR-1's function.

The authors have adequately addressed most of my previous concerns, however, I recommend minor revisions to further strengthen the study's rigor and interpretation:

Major suggestions

(1) This study relies heavily on aldicarb assays to support its conclusions. While these assays are valuable, their results may not fully align with direct assessment of neurotransmitter release from motor neurons. For instance, prior work has shown that two presynaptic modulators identified through aldicarb sensitivity assays exhibited no corresponding electrophysiological defects at the neuromuscular junction (Liu et al., J Neurosci 27: 10404-10413, 2007). Similarly, at least one study from the Kaplan lab has noted discrepancies between aldicarb assays and electrophysiological analyses. The authors should consider adding a few sentences in the Discussion to acknowledge this limitation and the potential caveats of using aldicarb assays, especially since some of the aldicarb assay results in this study are not easily interpretable.

Aldicarb assays have been used very successfully in identifying mutants with defects in chemical synaptic transmission, and entire genetic screens have been conducted this way. The reviewer is right, one needs to realize that it is the balance of excitation and inhibition at the NMJ of C. elegans, which underlies the effects on the rate of aldicarb-induced paralysis, not just cholinergic transmission. I.e. if a given mutant affects cholinergic and GABAergic transmission differently, things become difficult to interpret, particularly if also muscle physiology is affected. Therefore, we combined mutant analyses with cell-type specific rescue. We acknowledge that results are nonetheless difficult to interpret. We thus added a sentence in the first paragraph of the discussion.

(2) The manuscript states, "Elevated Ca2+ levels were not further enhanced in a flwr-1;mca-3 double mutant." (lines 549-550). However, Figure 7C does not include statistical comparisons between the single and double mutants of flwr-1 and mca-3. Please add the necessary statistical analysis to support this statement.

Because we only marked significant differences in that figure, and n.s. was not shown. This was stated in the figure legend.

(3) The term "Ca2+ influx" should be avoided, as this study does not provide direct evidence (e.g. voltage-clamp recordings of Ca2+ inward currents in motor neurons) for an effect of the flwr-1 mutation of Ca2+ influx. The observed increase in neuronal GCaMP signals in response to optogenetic activation of ChR2 may result from, or be influenced by, Ca2+ mobilization from of intracellular stores. For example, optogenetic stimulation could trigger ryanodine receptor-mediated Ca2+ release from the ER via calcium-induced calcium release (CICR) or depolarization-induced calcium release (DICR). It would be more appropriate to describe the observed increase in Ca2+ signal as "Ca2+ elevation" rather than increased "Ca2+ influx".

Ok, yes, we can do this, we referred by ‘influx’ to cytosolic Ca2+, that fluxes into the cytosol, be it from the internal stores or the extracellular. Extracellular influx, more or less, inevitably will trigger further influx from internal stores, to our understanding. We changed this to “elevated Ca2+ levels” or “Ca2+ level rise” or “Ca2+ level increase”.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

A thorough discussion on the impact of cell-autonomous versus non-cell-autonomous effects is necessary.

Revise and clarify the distinction between local and global Ca²⁺ changes.

see above.

Reviewer #2 (Recommendations for the authors):

Minor suggestions

(1) In "Few-Ubi was shown to facilitate recovery of neurons following intense synaptic activity (Yao et al.,....." (lines 283-284), please specify which aspects of neuronal recovery are influenced by the Flower protein.

We added “refilling of SV pools”.

(2) The abbreviation "Few-Ubi" is used for the Drosophila Flower protein (e.g., line 283, Figure 1A, and Figure 8A). Please clarify what "Ubi" stands for and verify whether its inclusion in the protein name is appropriate.

This is inconsistent across the literature, sometimes Fwe-Ubi is also referred to as FweA. We now added this term. Ubi refers to ubiquitous (“Therefore, we named this isoform fweubi because it is expressed ubiquitously in imaginal discs“) (Rhiner 2010)

(3) The manuscript uses "pflwr-1" (line 303 and elsewhere) to denote the flwr-1 promoter. This notation could be misleading, as it may be interpreted as a gene name. Please consider using either "flwr-1p" or "Pflwr-1" instead. Additionally, ensure proper italicization of gene names throughout the manuscript.

We changed this throughout. We will change to italicized at proof stage, it would be too timeconsuming to spot these incidents now.

(4) The authors tagged the C-terminus of FLWR-1 by GFP (lines 321). The fusion protein is referred to as "GFP::FLWR-1" throughout the manuscript. Please verify whether "FLWR-1::GFP" would be the more appropriate designation.

Thank you, yes, we changed this in the text, GFP is indeed N-terminal.

(5) In "This did not show any additive effects...." (line 363), please clarify what "This" refers to.

Altered to “The combined rescues did not show any additive effects…”

(6) In "..., supporting our previous finding of increased neurotransmitter release in GABAergic neurons" (lines 412-413), please provide a citation for the referenced previous study.

This refers to our aldicarb data within this paper, just further up in the text. We removed “previous”.

(7) Figure 4C, D examines the effect of flwr-1 mutation on body length in the genetic background of the unc-29 mutation, which selectively disrupts the levamisole-sensitive acetylcholine receptor. Please comment on the rationale for implicating only the levamisole receptor rather than the nicotinic acetylcholine receptor in muscle cells.

This was because we used a behavioral assay. Despite the fact that the homopentameric ACR16/N-AChR mediate about 2/3 of the peak currents in response to acute ACh application to the NMJ (e.g. Almedom et al., EMBO J, 2009), the acr-16 mutant has virtually no behavioral / locomotion phenotype. Likely, this is because the heteropentameric, UNC-29 containing LAChR, while only contributing 1/3 of the peak current, desensitizes much more slowly and thus unc-29 mutants show a severe behavioral phenotype (uncoordinated locomotion, etc.). We thus did not expect a major effect when performing the behavoral assay in acr-16 mutants and thus chose the unc-29 mutant background.

(8) In "we found no evidence ....insertion into the PM (Yao et al., 2009)", It appears that the cited paper was not authored by any of the current manuscript. Please confirm whether this citation is correctly attributed.

This sentence was arranged in a misleading way, we did not mean that we authored this paper. It was change in the text: “While a facilitating role of Flower in endocytosis appears to be conserved in C. elegans, in contrast to previous findings from Drosophila (Yao et al., 2009), we found no evidence that FLWR-1 conducts Ca2+ upon insertion into the PM.”

Changes from the earlier version to the revised version

Abstract

Abstract

The Flower protein was suggested to couple the fusion of synaptic vesicles (SVs) to their recycling in different model organisms. It is supposed to trigger activity-dependent bulk endocytosis by conducting Ca2+ at endocytic sites. However, this mode of action is debated. Here, we investigated the role of the Caenorhabditis elegans homolog FLWR-1 in neurotransmission. Our results confirm that FLWR-1 facilitates the recycling of SVs at the neuromuscular junction (NMJ). Ultrastructural analysis of synaptic boutons after hyperstimulation revealed an accumulation of large endocytic structures in flwr-1 mutants. These findings do not support a role of FLWR-1 in the formation of bulk endosomes but rather a function in their breakdown. Unexpectedly, loss of FLWR-1 led to increased neuronal Ca2+ levels in axon terminals during stimulation, particularly in GABAergic motor neurons, causing excitation-inhibition imbalance. We found that this increased NMJ transmission might be caused by deregulation of MCA-3, the nematode orthologue of the plasma membrane Ca2+ ATPase (PMCA). In vivo molecular interactions indicated that FLWR-1 may be a positive regulator of the PMCA and might influence its recycling through modification of plasma membrane levels of PI(4,5)P2.

Introduction

Removed in the revised version: Chemical synaptic transmission involves a cycle of biogenesis of synaptic vesicles (SVs), their fusion with the plasma membrane (PM), as well as their recycling by endocytosis and de novo formation in the endosome (Alabi and Tsien, 2012; Chanaday et al., 2019; Kononenko and Haucke, 2015; Rizzoli, 2014; Saheki and De Camilli, 2012). Coupling of SV exocytosis and endocytosis must be tightly controlled to avoid depletion of the reserve pool of SVs and to enable sustained neurotransmission (Haucke et al., 2011; Lou, 2018). Different hypotheses have been formulated as to how this is achieved within neurons. One hypothesis suggests that an SV-integral transmembrane protein called Flower could form ion channels that get inserted into the PM during SV fusion (Yao et al., 2009).

Earlier version: Subsequently, Flower may facilitate endocytosis by conducting Ca2+ into the cytoplasm, thus contributing to defining endocytic sites (Yao et al., 2017). In Drosophila melanogaster, Flower was proposed to increase phosphatidylinositol-4,5-bisphosphate (PI(4,5)P2) levels through Ca2+ microdomains, which was suggested to drive activity-dependent bulk endocytosis (ADBE) and formation of new SVs after prolonged, intense neurotransmission (Li et al., 2020). The hypothesis that Flower may have Ca2+ channel activity was further proposed based on sequence similarities between Flower and the Ca2+ selectivity filter of voltage-gated Ca2+ channels (VGCCs; Yao et al. (2009). Additional evidence from Drosophila suggests that Flower may regulate clathrin-mediated endocytosis (CME) in a Ca2+ independent fashion (Yao et al., 2017). However, while a facilitatory role of Flower in endocytosis appears to be evolutionarily conserved, and was observed in different organisms and tissues, including non-neuronal cells (Chang et al., 2018; Rudd et al., 2023; Yao et al., 2017), its Ca2+ channel activity and its influence on SV recycling is debated (Coelho and Moreno, 2020; Lou, 2018). The kinetics of Ca2+ influx mediated by Flower appear to be too slow and the amount conducted too low to have a major impact (Xue et al., 2012). Moreover, the function of Flower appears to depend on the cell type and the extent of synaptic activity (Chang et al., 2018; Yao et al., 2017). Apart from its role in endocytosis, Flower is involved in cell survival mechanisms during development (Coelho and Moreno, 2020; Costa-Rodrigues et al., 2021). Intriguingly, loss of the mammalian homolog of Flower can reduce tumor growth, suggesting an important function in tumor cell survival and indicating possible options for cancer treatment (Madan et al., 2019; Petrova et al., 2012). However, more research is needed to determine the exact signaling pathways by means of which Flower mediates cell survival (Costa-Rodrigues et al., 2021).

Revised version: Chemical synaptic transmission involves a cycle of biogenesis of synaptic vesicles (SVs), their fusion with the plasma membrane (PM), as well as their recycling by endocytosis and de novo formation in the endosome (Alabi and Tsien, 2012; Chanaday et al., 2019; Kononenko and Haucke, 2015; Rizzoli, 2014; Saheki and De Camilli, 2012). Coupling of SV exocytosis and endocytosis must be tightly controlled to avoid depletion of the reserve pool of SVs and to enable sustained neurotransmission (Haucke et al., 2011; Lou, 2018). Different hypotheses have been formulated as to how this is achieved within neurons. One hypothesis suggests that an SV-integral transmembrane protein called Flower could form ion channels that get inserted into the PM during SV fusion (Yao et al., 2009). Subsequently, Flower may facilitate endocytosis by conducting Ca2+ into the cytoplasm, thus contributing to defining endocytic sites (Yao et al., 2017). In Drosophila melanogaster, Flower was proposed to increase phosphatidylinositol-4,5-bisphosphate (PI(4,5)P2) levels through Ca2+ microdomains, which was suggested to drive activity-dependent bulk endocytosis (ADBE) and formation of new SVs after prolonged, intense neurotransmission (Li et al., 2020). The hypothesis that Flower may have Ca2+ channel activity was further proposed based on sequence similarities between Flower and the Ca2+ selectivity filter of voltage-gated Ca2+ channels (VGCCs; Yao et al. (2009). Additional evidence from Drosophila suggests that Flower may regulate clathrin-mediated endocytosis (CME) in a Ca2+ independent fashion (Yao et al., 2017). However, while a facilitatory role of Flower in endocytosis appears to be evolutionarily conserved, and was observed in different organisms and tissues, including non-neuronal cells (Chang et al., 2018; Rudd et al., 2023; Yao et al., 2017), its Ca2+ channel activity and its influence on SV recycling is debated (Coelho and Moreno, 2020; Lou, 2018). The kinetics of Ca2+ rise mediated by Flower appear to be too slow and the amount conducted too low to have a major impact (Xue et al., 2012). Moreover, the function of Flower appears to depend on the cell type and the extent of synaptic activity (Chang et al., 2018; Yao et al., 2017). Apart from its role in endocytosis, Flower is involved in cell survival mechanisms during development (Coelho and Moreno, 2020; Costa-Rodrigues et al., 2021). Intriguingly, loss of the mammalian homolog of Flower can reduce tumor growth, suggesting an important function in tumor cell survival and indicating possible options for cancer treatment (Madan et al., 2019; Petrova et al., 2012). However, more research is needed to determine the exact signaling pathways by means of which Flower mediates cell survival (Costa-Rodrigues et al., 2021).

Earlier version: Early studies of the Flower protein proposed a two-or three-transmembrane helical organization in which the C-terminus is exposed to the extracellular space and may thus mediate intercellular communication (Costa-Rodrigues et al., 2021; Rhiner et al., 2010). However, more recent research has shown that both N-and C-termini of Flower are likely cytosolic, and that the longest mammalian isoform consists of four transmembrane helices which are connected by short loops (Chang et al., 2018; Rudd et al., 2023). The genome of the nematode Caenorhabditis elegans is predicted to contain only a single isoform of Flower (FLWR-1; wormbase.org). C. elegans therefore may serve as an excellent model organism to further investigate the evolutionarily conserved role of Flower in neurotransmission and endocytosis. Here, we studied the function of C. elegans FLWR-1. We find that FLWR-1 localizes to SVs and to the PM and is involved in neurotransmission. We further show that FLWR-1 has a facilitatory but not essential role in endocytosis, confirming previous research in Drosophila and mammalian cells. Loss of FLWR-1 surprisingly conveys increased Ca2+ influx in optogenetically depolarized motor neurons. Yet, this is accompanied by reduced neurotransmitter release based on pharmacological assays and following optogentic stimulation, thus suggesting a deregulation of Ca2+ signaling in the presynapse. This is associated with defective SV recycling at the level of the endosome and thus reduced SV numbers upon sustained stimulation. The increased Ca2+ influx is more pronounced in γ-aminobutyric acid (GABA) releasing neurons and leads to an excitation-inhibition (E/I) imbalance at the neuromuscular junction (NMJ) through increased release of the neurotransmitter GABA. A function of FLWR-1 in endocytosis appears to affect also the plasma membrane Ca2+-ATPase MCA-3, required for extrusion of Ca2+ from the cytosol. This may explain the increased Ca2+ levels during stimulation in flwr-1 mutant neurons. Last, our findings suggest a possible direct molecular interaction between FLWR-1 and MCA-3.

Revised version: Early studies of the Flower protein proposed a two- or three-transmembrane helical organization in which the C-terminus is exposed to the extracellular space and may thus mediate intercellular communication (Costa-Rodrigues et al., 2021; Rhiner et al., 2010). However, more recent research has shown that both N- and C-termini of Flower are likely cytosolic, and that the longest mammalian isoform consists of four transmembrane helices which are connected by short loops (Chang et al., 2018; Rudd et al., 2023). The genome of the nematode Caenorhabditis elegans is predicted to contain only a single isoform of Flower (FLWR-1; wormbase.org). C. elegans therefore may serve as an excellent model organism to further investigate the evolutionarily conserved role of Flower in neurotransmission and endocytosis. Here, we studied the function of C. elegans FLWR-1. We find that FLWR-1 localizes to SVs and to the PM and is involved in neurotransmission. We further show that FLWR-1 has a facilitatory but not essential role in endocytosis, confirming previous research in Drosophila and mammalian cells. Loss of FLWR-1 surprisingly conveys increased Ca2+ levels in optogenetically depolarized motor neurons. Yet, this is accompanied by reduced neurotransmitter release based on pharmacological assays and following optogentic stimulation, thus suggesting a deregulation of Ca2+ signaling in the presynapse. This is associated with defective SV recycling at the level of the endosome and thus reduced SV numbers upon sustained stimulation. The increased Ca2+ level is more pronounced in γ-aminobutyric acid (GABA) releasing neurons and leads to an excitation-inhibition (E/I) imbalance at the neuromuscular junction (NMJ) through increased release of the neurotransmitter GABA. A function of FLWR-1 in endocytosis appears to affect also the plasma membrane Ca2+-ATPase MCA-3, required for extrusion of Ca2+ from the cytosol. This may explain the increased Ca2+ levels during stimulation in flwr-1 mutant neurons. Last, our findings suggest a possible direct molecular interaction between FLWR-1 and MCA-3.

Materials and methods

Molecular biology

For Plasmids used in this paper and details of their generation, see Table S1. pZIM902 [punc-17b::GCaMP6fOpt] was a gift from Manuel Zimmer. p1676 [punc-17(short)::TagRFP::ELKS-1] was kindly provided by the lab of Zhao-Wen Wang. pJH2523 [punc-25::GCaMP3::UrSL2::wCherry] was a gift from Mei Zhen (Addgene plasmid # 191358; http://n2t.net/addgene:191358; RRID:Addgene_191358) (Lu et al., 2022).

Cultivation of C. elegans

Animals were kept at 20 °C on nematode growth medium (NGM) plates seeded with OP50-1 bacteria (Brenner, 1974). For optogenetic experiments, OP50-1 was supplemented with 200 µM all-trans retinal (ATR) prior to seeding and animals were kept in darkness. Transgenic animals carrying extrachromosomal arrays were generated by microinjection into the gonads (Fire, 1986). L4 staged larvae were picked ∼18 h before experiments and tested on at least three separate days with animals picked from different populations. Transgenic animals were selected by fluorescent markers using a Leica MZ16F dissection stereo microscope. Integration of the sybIs8965 transgenic array (PHX8965 strain) was performed by SunyBiotech. The RB2305 strain containing the flwr-1(ok3128) allele was outcrossed three times to N2 wild type. For an overview of strains, genotypes and transgenes, see Table S2.

Counting live progeny

To count living progeny per animal, five L4 larvae were singled onto NGM plates seeded with 50 µl OP50-1. Animals were then picked onto fresh plates after two days and on each of the following two days and removed on day five. Living progeny per animal were counted after reaching adulthood and summed up over the three plates on which each parental animal laid eggs, respectively. Experiments were performed blinded to the genotype.

Pharmacological assays

1.5 mM aldicarb and 0.25 mM levamisole plates were prepared by adding the compounds to liquid NGM prior to plate pouring (Mahoney et al., 2006). Aldicarb (Sigma Aldrich, USA) was kept as a 100 mM stock solution in 70 % ethanol. Levamisole (Sigma Aldrich, USA) was stored as a 200 mM solution dissolved in ddH2O. 12 - 30 young adult animals were transferred to each plate and tested every 15 (levamisole) or 30 (aldicarb) minutes. Assays were performed blinded to the genotype and control groups (wild type and flwr-1 mutants), measured in parallel on the same day. Animals that did not respond after being prodded three times with a hair pick were counted as paralyzed. Worms that crawled of the plate were disregarded from analysis.

Measurement of swimming cycles and crawling speed using the multi-worm-tracker (MWT)

Swimming cycles were measured as described previously (Vettkotter et al., 2022). In short, worms were washed three times with M9 buffer to remove OP50 and transferred onto 3.5 cm NGM plates with 800 µl M9. Animals were visualized using the multiworm tracker (MWT) platform (Swierczek et al., 2011) equipped with a Falcon 4M30 camera (DALSA). Channelrhodopsin-2 (ChR2) was stimulated with 470 nm light at 1 mW/mm2 intensity for 90 s. 30 s videos were captured and thrashing was analyzed using the “wrmTrck” plugin for ImageJ (Nussbaum-Krammer et al., 2015). The automatically generated tracks were validated using a custom written Python script (https://github.com/dvettkoe/SwimmingTracksProcessing). Crawling speed was measured with the same setup and as described before (Vettkotter et al., 2022) yet animals were transferred to empty, unseeded NGM plates after washing. Prior to the measurement, animals were incubated in darkness for 15 min which is necessary for a switch from local to global search for food and thus a constant crawling speed (Calhoun et al., 2014). The crawling speed was recorded with the “Multi-Worm Tracker” software and extracted using the “Choreography” software (Swierczek et al., 2011). A custom-written Python script was used to summarize the measured tracks (https://github.com/dvettkoe/MWT_Analysis).

For manual counting of swimming cycles (Figure 9 – Figure supplement 1), animals were transferred on an NGM plate in M9 buffer, and filmed with a Canon Powershot G11 camera, for 30 sec. Swimming cycles were counted manually during replay of the video.

Measurement of body length

Body length was measured as described previously (Liewald et al., 2008; Seidenthal et al., 2022). In short, single animals were transferred onto unseeded NGM plates and illuminated with light from a 50 W HBO lamp which was filtered with a 450 - 490 nm bandpass excitation filter. ChR2 was stimulated with 100 µW/mm2 light intensity. A 665 - 715 nm filter was used to avoid unwanted activation by brightfield light. Body length was analyzed using the “WormRuler” software (version 1.3.0) and normalized for the average skeleton length before stimulation (Seidenthal et al., 2022). Values deviating more than 20 % from the initial body length were discarded as they result from artifacts in the background correction (and are biomechanically impossible). Basal body length was calculated from the skeleton length in pixels which is generated by the WormRuler software. This was converted to mm and averaged for each animal over a duration of five seconds.

Light microscopy and fluorescence quantification

Removed in the revised version: Animals were placed upon 7 % agarose pads in M9 buffer. Worms were immobilized by either using 20 mM levamisole in M9 or Polybead polystyrene microspheres (Polysciences) for experiments involving GABAergic stimulation. For optogenetic experiments, single animals were placed upon pads to avoid unwanted pre-activation of channelrhodopsins.

Earlier version: Imaging was performed using a Axio Observer Z1 microscope (Zeiss, Germany). Proteins were excited using 460 nm and 590 nm LEDs (Lumen 100, Prior Scientific, UK) coupled via a beamsplitter. Background correction was conducted by placing a region of interest (ROI) within the animal yet avoiding autofluorescence.

Revised version: Animals were placed upon 7 % agarose pads in M9 buffer. Worms were immobilized by either using 20 mM levamisole in M9 or Polybead polystyrene microspheres (Polysciences) for experiments involving GABAergic stimulation. For optogenetic experiments, single animals were placed upon pads to avoid unwanted pre-activation of channelrhodopsins. Imaging was performed using a Axio Observer Z1 microscope (Zeiss, Germany). Proteins were excited using 460 nm and 590 nm LEDs (Lumen 100, Prior Scientific, UK) coupled via a beamsplitter. Background correction was conducted by placing a region of interest (ROI) within the animal yet avoiding autofluorescence.

GCaMP and pHluorin imaging and ChrimsonSA stimulation was performed using a 605 nm beamsplitter (AHF Analysentechnik, Germany) which was combined with a double band pass filter (460 - 500 and 570 - 600 nm). A single bandpass emission filter was used (502.5 - 537.5 nm). pOpsicle assays were performed as described previously using a 100x objective (Seidenthal et al., 2023). We extended the ChrimsonSA stimulation light pulse to 30 s to be consistent with the stimulus length of electron microscopy (EM) and electrophysiology experiments. Moreover, 2 x 2 binning was applied. Image sequences with 200 ms light exposure (5 frames per second) were acquired with an sCMOS camera (Kinetix 22, Teledyne Photometrics, USA). Video acquisition was controlled using the µManager v.1.4.22 software (Edelstein et al., 2014). The timing of LED activation was managed using a Autohotkey script. Dorsal nerve cord (DNC) fluorescence was quantified with ImageJ by placing a ROI with the Segmented Line tool. XY-drift was corrected using the Template Matching plugin if necessary. Animals showing excessive z-drift were discarded. A custom written python script was used to summarize background subtraction and normalization to the average fluorescence before pulse start (https://github.com/MariusSeidenthal/pHluorin_Imaging_Analysis).

GCaMP imaging was conducted similarly with the same setup and light intensity for ChrimsonSA stimulation (40 µW/mm2) yet without binning. GCaMP fluorescence was quantified by placing ROIs around four synaptic puncta. Puncta fluorescence was averaged prior to background correction and normalization. Moreover, the bleach correction function of the python script was used which corrects fluorescence traces with strong bleaching. Strong bleaching is defined as the fluorescence intensity of the mean of the last second of measurement being less than 85% of the first second. The script performs linear regression analysis for DNC fluorescent traces and background ROI and corrects values prior to background correction. The increase in GCaMP fluorescence during stimulation was highly dependent on basal fluorescence values before stimulation and transient Ca2+ signals. This caused some measurements to show very strong increases in fluorescence. To avoid distortion of statistical analyses, we performed outlier detection for the increase during stimulation on all datasets with the Graphpad Prism 9.4.1 Iterative Grubb’s method. The alpha value (false discovery rate) was set to 0.01.

For pOpsicle assays, animals were assessed for whether they exhibited a significant increase in fluorescence during stimulation, which was necessary for analysis of fluorescence decay after stimulation. A strong signal was determined as the maximum background corrected fluorescence during the light pulse (moving average of 5 frames) being larger than the average before stimulation plus 3x the standard deviation of the background corrected fluorescence before stimulation. Regression analysis of pHluorin fluorescence decay was performed as described previously by using Graphpad Prism 9.4.1 and fitting a “Plateau followed one-phase exponential decay” fit beginning with the first time point after stimulation to single measurements (Seidenthal et al., 2023). Animals that displayed no increase during stimulation as well as animals showing no decay or spontaneous signals after stimulation were discarded from analysis.

Comparison of dorsal and ventral (VNC) nerve cord fluorescence of GFP::FLWR-1 and mCherry::SNB-1 was conducted by imaging the posterior part of the animal where an abundance of synapses can be found. Images were acquired using the same beamsplitter and excitation filter as used for pOpsicle experiments but equipped with a 500 – 540 nm/600 – 665 nm double bandpass emission filter (AHF Analysentechnik, Germany) and 2 x 2 binning. The Arduino script “AOTFcontroller” was used to control synchronized 2-color illumination (Aoki et al., 2024). VNC and DNC fluorescence were analyzed by choosing single frames from acquired z-stacks in which nerve cords were well focused. Fluorescence intensities were quantified by placing a Segmented Line ROI. Kymographs were generated with the ImageJ Multi Kymograph function.

Confocal laser scanning microscopy was performed on a LSM 780 microscope (Zeiss, Germany) equipped with a Plan-Apochromat 63-x oil objective. The Zeiss Zen (blue edition) Tile Scan and Z-Stack functions were used to generate overview images of animals expressing GFP::FLWR-1. Images were processed and maximum z-projection performed using ImageJ. Colocalization analysis was performed by placing a Segmented Line ROI onto the DNC and plotting the fluorescence profile along the selected ROI for both color channels. Correlation of fluorescence intensities was then conducted using the GraphPad Prism 9.4.1 Pearson correlation function.

Fluorescence of GFP within coelomocytes (CCs) was visualized with a 100-x objective. An eGFP filter cube (AHF Analysentechnik, Germany) was used and fluorescence was excited with the 460 nm LED. Images were acquired using 2 x 2 binning and 50 ms exposure for ssGFP and 200 ms for GFP fused to the PH-domain. Endocytosed ssGFP fluorescence was quantified by drawing a ROI around all CCs found within any worm imaged using the Freehand selections tool and performing background correction. CC fluorescence was then normalized to the average fluorescence in wild type animals on the respective measurement day and CC location (anterior/middle/posterior). This was necessary since anterior CCs generally had a higher fluorescence than those in other locations. Membrane localization of PH-domain GFP was determined by drawing a Segmented Line ROI around the perimeter of the cell and drawing a rough outline of the cell interior using the Freehand selections tool. Membrane fluorescence was then corrected by subtracting the measured interior fluorescence for each cell.

In BiFC experiments, split mVenus was imaged with a 100x objective, the eGFP filter cube and 500 ms exposure. Fluorescence was excited with 460 nm.

MCA-3b::YFP was imaged with a 40-x objective, the eGFP filter cube and 200 ms exposure. Fluorescence was excited with 460 nm. Synaptic puncta towards the posterior part of the animal were imaged and quantified by drawing a Freehand selections ROI around all visible and well-focused synaptic puncta in an animal before averaging. Similarly, interpuncta fluorescence was quantified by drawing a ROI around several locations inbetween the examined puncta.

Primary neuronal cell culture

Primary neuronal cell cultures were generated from embryos of SNG-1::pHluorin expressing animals as described before (Seidenthal et al., 2023) yet seeded into perfusion chambers (0.4 mm channel height, ibidi). pHluorin fluorescence was excited with 460 nm and imaged using the eGFP filter cube and 1000 ms exposure time. The buffers used to quench surface and dequench intracellular pHluorin were adapted from Dittman and Kaplan (2006). Cells were washed three times with the respective buffer prior to imaging by pipetting 1 ml into the reservoir wells of the perfusion chamber and removing the same amount of liquid from the other side. Individual cells were analyzed by placing a Segmented Line ROI onto extending neurites and performing background correction. The fraction of SNG-1::pHluorin which resides in SVs was then estimated by dividing the fluorescence originating from vesicular pHluorin by the total fluorescence originating in unquenched pHluorin: F0 → the basal neurite fluorescence in control saline buffer

FNH4Cl → neurite fluorescence during addition of a buffer containing NH4Cl which unquenches intravesicular pHluorin

FpH=5.6 → neurite fluorescence with surface quenched pHluorin by low pH

Electrophysiology

Electrophysiological recordings of body wall muscle cells (BWMs) were performed in dissected adult worms as previously described (Liewald et al., 2008). Animals were immobilized with Histoacryl L glue (B. Braun Surgical, Spain) and a lateral incision was made to access NMJs along the anterior VNC. The basement membrane overlying BWMs was enzymatically removed by 0.5 mg/ml collagenase for 10 s (C5138, Sigma-Aldrich, Germany). Integrity of BWMs and nerve cord was visually examined via DIC microscopy.

Recordings from BWMs were acquired in whole-cell patch-clamp mode at 20-22 °C using an EPC-10 amplifier equipped with Patchmaster software (HEKA, Germany). The head stage was connected to a standard HEKA pipette holder for fire-polished borosilicate pipettes (1B100F-4, Worcester Polytechnic Institute, USA) of 4–10 MΩ resistance. The extracellular bath solution consisted of 150 mM NaCl, 5 mM KCl, 5 mM CaCl2,1 mM MgCl2,10 mM glucose, 5 mM sucrose, and 15 mM HEPES, pH 7.3, with NaOH, ∼330 mOsm. The internal/patch pipette solution consisted of 115 mM K-gluconate, 25 mM KCl, 0.1 mM CaCl2, 5 mM MgCl2, 1 mM BAPTA, 10 mM HEPES, 5 mM Na2ATP, 0.5 mM Na2GTP, 0.5 mM cAMP, and 0.5 mM cGMP, pH 7.2, with KOH, ∼320 mOsm.

Voltage clamp experiments were conducted at a holding potential of −60 mV. Light activation was performed using an LED lamp (KSL-70, Rapp OptoElectronic, Hamburg, Germany; 470 nm, 8 mW/mm²) and controlled by the Patchmaster software. Subsequent analysis was performed using Patchmaster and Origin (Originlabs). Analysis of mPSCs was conducted with MiniAnalysis (Synaptosoft, Decatur, GA, USA, version 6.0.7) and rate and amplitude of mPSCs were analyzed in 1 s bins. Exponential decay of mPSCs during stimulation was calculated with Graphpad Prism 9.4.1 by fitting a one-phase exponential decay beginning with the first time point during stimulation.

Transmission electron microscopy

Prior to high-pressure freezing (HPF), L4 animals were transferred to freshly seeded E. coli OP50-1 dishes supplemented with or without 0.1 mM ATR. HPF fixation was performed on young adult animals as described previously (Kittelmann et al., 2013b; Weimer, 2006). Briefly, 20 - 40 animals were transferred into a 100 µm deep aluminum planchette (Microscopy Services) filled with E. coli (supplemented with or without ATR, respectively), covered with a sapphire disk (0.16 mm) and a spacer ring (0.4 mm; engineering office M. Wohlwend) for photostimulation. To prevent pre-activation of ChR2, all preparations were carried out under red light. Animals were continuously illuminated for 30 s with a laser (470 nm, ∼20 mW/mm²) followed 5 s later by HPF at −180 °C under 2100 bar pressure in an HPM100 (Leica Microsystems). Frozen specimens were transferred under liquid nitrogen into a Reichert AFS machine (Leica Microsystems) for freeze substitution.

Samples were incubated with tannic acid (0.1% in dry acetone) fixative at −90 °C for 100 h. Afterward, a process of washing was performed for substitution with acetone, followed by incubation of the samples in 2% OsO4 (in dry acetone) for 39.5 h while the temperature was slowly increased up to room temperature. Subsequently, samples were embedded in epoxy resin (Agar Scientific, AGAR 100 Premix kit – hard) by increasing epoxy resin concentrations from 50% to 90% at room temperature and 100% at 60 °C for 48 h.

Electron micrographs of 2 – 5 individual animals, and from 12 – 22 synapses per treatment, were acquired by cutting cross sections at a thickness of 40 nm, transferring cross sections to formvar- or pioloform covered copper slot grids. Specimens were counterstained in 2.5% aqueous uranyl acetate for 4 min, followed by washing with distilled water, and incubation in Reynolds lead citrate solution for 2 min in a CO2-free chamber with subsequent washing steps in distilled water.

VNC regions were then imaged with a Zeiss 900 TEM, operated at 80 kV, with a Troendle 2K camera. Images were scored and tagged blind in ImageJ (version 1.53c, National Institute of Health) as described previously, and analyzed using SynapsEM (Vettkotter et al., 2022; Watanabe et al., 2020). Since the number of synaptic organelles varied between synapses of different sizes, their counts were normalized to the average synaptic profile area of 145,253 nm2. Large vesicles (LVs) are defined by their size larger than 50 nm and empty lumen as described by (Kittelmann et al., 2013b), endosomes are defined as structures larger than 100 nm and located more than 50 nm away from the dense projection, as described by (Watanabe et al., 2014); however, by an alternative definition of endosome, structures with sizes larger than SVs and DCVs, with an electron dense lumen were also included (Kittelmann et al., 2013b).

Protein alignment and structure prediction

Protein alignment was performed with ClustalX (Larkin et al., 2007) and shading of conserved residues created with Boxshade (https://junli.netlify.app/apps/boxshade/). Transmembrane domains as well as membrane orientation were predicted by entering the predicted amino acid sequence to DeepTMHMM (Hallgren et al., 2022). FLWR-1 structure was modeled and visualized with Alphafold3 (AF3) (Abramson et al., 2024). FLWR-1 tetramers have been modeled using AlphafoldMultimer or AF3 (Evans et al., 2022). Protein structures were analyzed and edited using PyMol v2.5.2. Docking of PI(4,5)P2 to the FLWR-1 AF3 model was conducted in PyMol using the DockingPie plugin running the Autodock/Vina analysis (Eberhardt et al., 2021; Rosignoli and Paiardini, 2022).

Statistical Analysis and generation of graphs

Graphs were created and statistical analyses were performed using GraphPad Prism 9.4.1. Data was displayed as mean ± standard error of the mean (SEM) unless noted otherwise. Unpaired t-tests were used to compare two datasets and one-way ANOVAs or two-way ANOVAs if three or more datasets were assessed, according to the number of tested conditions. Fitting of a mixed-effects model was used instead of two-way ANOVA if the number of data points was not constant between different datasets for example if the number of animals differed between time points. Mann-Whitney tests were performed if two datasets were compared that were not normally distributed and datasets were depicted as median with interquartile range (IQR) in statistical analyses. Kruskal-Wallis tests were conducted if not normally distributed datasets were assessed. Iterative Grubb’s outlier detection was used when required and as indicated. The alpha value (false discovery rate) was set to 0.01. Representations of the exon-intron structures of genes were created using the “Exon-Intron Graphic Maker” (Bhatla, 2012).

Results

FLWR-1 is involved in neurotransmission

Earlier version: The amino acid sequence of C. elegans FLWR-1 is conserved with its D. melanogaster (Fwe-Ubi) and Homo sapiens (hFwe4) homologues (Fig. 1A). Fwe-Ubi was shown to facilitate recovery of neurons following intense synaptic activity (Yao et al., 2009; Yao et al., 2017). To investigate whether this involvement of Flower in neurotransmission is evolutionarily conserved, we studied a mutant lacking most of the flwr-1 coding region (ok3128, Fig. 1B) (Consortium, 2012). Animals lacking FLWR-1 did not display severe phenotypes. Basal locomotion in liquid (Fig. 1C, D seconds 0 – 30) and body length of young adults (Figure 1 – Figure supplement 1A) were not significantly different from the respective values of wild type animals. However, fertility appeared slightly reduced (Figure 1 – Figure supplement 1B). Basal crawling speed, on average, was not different between wild type, flwr-1 mutants and rescued animals (genomic sequence and 2kB promoter; Figure 1 – Figure supplement 1C).

Revised version: The amino acid sequence of C. elegans FLWR-1 is conserved with its D. melanogaster (Fwe-Ubi/FweA) and Homo sapiens (hFwe4) homologues (Fig. 1A). Fwe-Ubi/FweA was shown to facilitate recovery of neurons, i.e. refilling of SV pools, following intense synaptic activity (Yao et al., 2009; Yao et al., 2017). To investigate whether this involvement of Flower in neurotransmission is evolutionarily conserved, we studied a mutant lacking most of the flwr-1 coding region (ok3128, Fig. 1B) (Consortium, 2012). Animals lacking FLWR-1 did not display severe phenotypes. Basal locomotion in liquid (Fig. 1C, D seconds 0 – 30) and body length of young adults (Figure 1 – Figure supplement 1A) were not significantly different from the respective values of wild type animals. However, fertility appeared slightly reduced (Figure 1 – Figure supplement 1B). Basal crawling speed, on average, was not different between wild type, flwr-1 mutants and rescued animals (genomic sequence and 2kB promoter; Figure 1 – Figure supplement 1C).

Loss of FLWR-1 induces defects in neurotransmission following intense stimulation

Earlier version: (A) Alignment of the amino acid sequences of FLWR-1 to hFwe4 (Homo sapiens) and Fwe-Ubi (Drosophila melanogaster). Shading depicts evolutionary conservation of amino acid residues (identity – black; homology – grey). Position of TM helices indicated in red refers to the FLWR-1 sequence. (B) Schematic representation of the flwr-1/F20D1.1 gene locus and the size of the ok3128 deletion. Bars represent exons and connecting lines introns. (C) Mean (± SEM) swimming cycles of animals expressing ChR2(H134R) in cholinergic motor neurons (unc-17 promoter). All animals were treated with ATR. A 90 s light pulse (470 nm, 1 mW/mm2) was applied after 30 s as indicated by the blue shade. Number of animals accumulated from N = 3 biological replicates: wild type = 80 – 88, flwr-1 = 80 – 91, FLWR-1 rescue = 62 – 75. (D) Statistical analysis of swimming speed at different time points as depicted in (C). Mean (± SEM). Each dot represents a single animal. Mixed-effects model analysis with Tukey’s correction. Only statistically significant differences are depicted. **p < 0.01, *** p < 0.001. (E) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb. N = 4 biological replicates. Two-way ANOVA with Tukey’s correction. ns not significant p > 0.05, *** p < 0.001.

Revised version: (A) Alignment of the amino acid sequences of FLWR-1 to hFwe4 (Homo sapiens) and Fwe-Ubi/FweA (Drosophila melanogaster). Shading depicts evolutionary conservation of amino acid residues (identity – black; homology – grey). Position of TM helices indicated in red refers to the FLWR-1 sequence. (B) Schematic representation of the flwr-1/F20D1.1 gene locus and the size of the ok3128 deletion. Bars represent exons and connecting lines introns. (C) Mean (± SEM) swimming cycles of animals expressing ChR2(H134R) in cholinergic motor neurons (unc-17 promoter). All animals were treated with ATR. A 90 s light pulse (470 nm, 1 mW/mm2) was applied after 30 s as indicated by the blue shade. Number of animals accumulated from N = 3 biological replicates: wild type = 80 – 88, flwr-1 = 80 – 91, FLWR-1 rescue = 62 – 75. (D) Statistical analysis of swimming speed at different time points as depicted in (C). Mean (± SEM). Each dot represents a single animal. Mixed-effects model analysis with Tukey’s correction. Only statistically significant differences are depicted. **p < 0.01, *** p < 0.001. (E) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb. N = 4 biological replicates. Two-way ANOVA with Tukey’s correction. ns not significant p > 0.05, *** p < 0.001.

Removed in the revised version: Loss of FLWR-1 does not change body length or locomotion speed but reduces the number of living progeny.

Removed in the revised version: (A) Mean (± SEM) body length of wild type and flwr-1(ok3128) mutants in mm. Unpaired t-test. ns not significant. Number of animals accumulated from N = 4 biological replicates: wild type = 29, flwr-1 = 27. (B) Mean (± SEM) number of living progeny per animal. Unpaired t-test. ** p < 0.01. Number of animals accumulated from N = 4 biological replicates: wild type = 20, flwr-1 = 18. (C) Mean (± SEM) crawling speed. N = 2 biological replicates.

Removed in the revised version: FLWR-1 is predicted to be a tetraspan transmembrane protein and is transported by UNC-104 kinesin.

Removed in the revised version: (A) Graphical representation of the results of DeepTMHMM prediction of membrane orientation of the FLWR-1 protein based on its amino acid sequence. The red shapes indicate the presence of four transmembrane domains. (B) Alphafold3 prediction of FLWR-1 protein structure. The coloring represents the calculated predicted local distance difference test (plDDT) as shown below the structure. Higher plDDT values indicate a higher confidence of correct prediction. (C) Heatmap plot depicting flwr-1(F20D1.1) single-cell RNAseq data generated by the CeNGEN database. Coloring indicates transcripts per million (TPM) per tissue normalized to the average expression as indicated in the legend. The size of the dots represents the percentage of cells of this cell type expressing the gene. (D) Mean (± SEM) swimming cycles of animals expressing ChR2(H134R) in cholinergic motor neurons (unc-17 promoter). All animals were treated with ATR. A 90 s light pulse (470 nm, 1 mW/mm2) was applied after 30 s as indicated by the blue shade. Number of animals accumulated from N = 2 biological replicates: wild type = 63 – 99, flwr-1 = 48 – 66, GFP::FLWR-1 rescue = 37 – 52. (E) Statistical analysis of swimming speed at different time points as depicted in (D). Mean (± SEM). Each dot represents a single animal. Mixed-effects model analysis with Tukey’s correction. Only statistically significant differences are depicted. *** p < 0.001. (F) Example images depicting GFP::FLWR-1 fluorescence in nerve ring and nerve cords in wild type and unc-104(e1265) mutants. Scale bar, 20 µm.

Removed in the revised version: FLWR-1 expression in body wall muscle cells partially rescues aldicarb resistance of flwr-1 mutants, but aldicarb resistance is not affected through AChRs. (A) Mean (± SEM) fraction of moving animals after exposure to 0.25 mM levamisole. N = 3 biological replicates. Two-way ANOVAs with Tukey’s correction. No significant differences were found. (B) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb. FLWR-1 is optionally expressed in BWMs (pmyo-3) only or combined with expression in cholinergic (punc-17) or GABAergic (punc-47) neurons in a flwr-1(ok3128) mutant background. N = 3 - 8 biological replicates. Two-way ANOVAs with Tukey’s correction. Only non-significant differences are depicted. For all other comparisons p < 0.001. (C) Mean (± SEM) fraction of moving animals after exposure to 0.5 mM aldicarb, compared in wild type, flwr-1(ok3128), acr-16(ok789) and flwr-1; acr-16 double mutants. N = 3 biological replicates. Two-way ANOVAs with Tukey’s correction. (D) Loss of GABAergic transmission conveys aldicarb hypersensitivity in flwr-1 mutants. Mean (± SEM) fraction of moving animals after exposure to 0.1 mM aldicarb. N = 3 biological replicates.

Earlier version: These findings indicate that FLWR-1 is unlikely to have an essential function in neurotransmission, at least not during basal in vivo activity, but rather a regulatory or facilitatory one. Previously we showed that mutations which only weakly affect basal locomotion can severely affect recovery of swimming speed after strong optogenetic stimulation of cholinergic neurons (Yu et al., 2018). We therefore subjected flwr-1(ok3128) mutants expressing channelrhodopsin-2 (ChR2; variant H134R), which were treated with the chromophore all-trans retinal (ATR), to blue light during swimming (Liewald et al., 2008; Nagel et al., 2005). Photostimulation resulted in a stop of all swimming during the light pulse, followed by a slow recovery in the dark. Indeed, loss of FLWR-1 led to a significantly slowed recovery of swimming locomotion (Fig. 1C, D). This could be fully rescued by transgenic expression of genomic flwr-1 including a two kilobase sequence upstream of the putative start codon, hereafter called pflwr-1 (as the promoter of flwr-1). To further evaluate the involvement of FLWR-1 in neurotransmission, we exposed flwr-1 mutants to aldicarb. This acetylcholine esterase (AChE) inhibitor induces paralysis due to accumulation of ACh in the synaptic cleft (Blazie and Jin, 2018; Mahoney et al., 2006). Resistance to aldicarb indicates either reduced release or detection of ACh or, alternatively, increased inhibitory signaling (Janosi et al., 2024; Vashlishan et al., 2008). Loss of FLWR-1 led to a significant delay in paralysis indicating an involvement in transmission at the NMJ (Fig. 1E). Full rescue of the flwr-1 mutant phenotypes suggests that expression from the 2 kB fragment of the endogenous promoter fully recapitulates the native expression in the context of NMJ function.

Revised version: These findings indicate that FLWR-1 is unlikely to have an essential function in neurotransmission, at least not during basal in vivo activity, but rather a regulatory or facilitatory one. Previously we showed that mutations which only weakly affect basal locomotion can severely affect recovery of swimming speed after strong optogenetic stimulation of cholinergic neurons (Yu et al., 2018). We therefore subjected flwr-1(ok3128) mutants expressing channelrhodopsin-2 (ChR2; variant H134R), which were treated with the chromophore all-trans retinal (ATR), to blue light during swimming (Liewald et al., 2008; Nagel et al., 2005). Photostimulation resulted in a stop of all swimming during the light pulse, followed by a slow recovery in the dark. Indeed, loss of FLWR-1 led to a significantly slowed recovery of swimming locomotion (Fig. 1C, D). This could be fully rescued by transgenic expression of genomic flwr-1 including a two kilobase sequence upstream of the putative start codon, hereafter called flwr-1p (as the promoter of flwr-1). To further evaluate the involvement of FLWR-1 in neurotransmission, we exposed flwr-1 mutants to aldicarb. This acetylcholine esterase (AChE) inhibitor induces paralysis due to accumulation of ACh in the synaptic cleft (Blazie and Jin, 2018; Mahoney et al., 2006). Resistance to aldicarb indicates either reduced release or detection of ACh or, alternatively, increased inhibitory signaling (Janosi et al., 2024; Vashlishan et al., 2008). Loss of FLWR-1 led to a significant delay in paralysis indicating an involvement in transmission at the NMJ (Fig. 1E). Full rescue of the flwr-1 mutant phenotypes suggests that expression from the 2 kB fragment of the endogenous promoter fully recapitulates the native expression in the context of NMJ function.

FLWR-1 is expressed in excitable cells and localizes to synaptic vesicles

Earlier version: Invertebrate and vertebrate homologues of FLWR-1 are predicted to be membrane proteins containing four transmembrane helices with both C-and N-termini located in the cytosol (Chang et al., 2018; Rudd et al., 2023; Yao et al., 2009). In silico predictions using the FLWR-1 sequence of 166 amino acids suggest evolutionary conservation of the tetraspan structure (Figure 1 – Figure supplement 2A) (Hallgren et al., 2022). Accordingly, Alphafold3 predicts a protein structure of FLWR-1 that implies a four helical configuration, with helix lengths that could span biological membranes, and three additional, shorter α-helices (Figure 1 – Figure supplement 2B) (Abramson et al., 2024). We sought to determine the cellular as well as subcellular localizations of FLWR-1 by tagging its C-terminus with GFP and expressing the construct using the endogenous promoter (Fig. 2A - C). Green fluorescence could be observed in developing embryos in the uterus, as well as in various tissues including neurons, body wall and pharyngeal muscles (Fig. 2A). FLWR-1 localized to neurites and cell bodies of nerve ring neurons, representing the central nervous system of the nematode (Fig. 2B) (Ward et al., 1975). These results are in agreement with single-cell RNA-sequencing data obtained from C. elegans which show near ubiquitous expression of flwr-1/F20D1.1 (Figure 1 – Figure supplement 2C, Table S3; Taylor et al. (2021)). The GFP::FLWR-1 fusion protein was functional as demonstrated by rescue of the swimming phenotype (recovery from cholinergic neuron photostimulation) observed in the flwr-1 mutant background (Figure 1 – Figure supplement 2D, E).

Revised version: Invertebrate and vertebrate homologues of FLWR-1 are predicted to be membrane proteins containing four transmembrane helices with both C- and N-termini located in the cytosol (Chang et al., 2018; Rudd et al., 2023; Yao et al., 2009). In silico predictions using the FLWR-1 sequence of 166 amino acids suggest evolutionary conservation of the tetraspan structure (Figure 1 – Figure supplement 2A) (Hallgren et al., 2022). Accordingly, Alphafold3 predicts a protein structure of FLWR-1 that implies a four helical configuration, with helix lengths that could span biological membranes, and three additional, shorter α-helices (Figure 1 – Figure supplement 2B) (Abramson et al., 2024). We sought to determine the cellular as well as subcellular localizations of FLWR-1 by tagging its N-terminus with GFP and expressing the construct using the endogenous promoter (Fig. 2A – C). Green fluorescence could be observed in developing embryos in the uterus, as well as in various tissues including neurons, body wall and pharyngeal muscles (Fig. 2A). FLWR-1 localized to neurites and cell bodies of nerve ring neurons, representing the central nervous system of the nematode (Fig. 2B) (Ward et al., 1975). These results are in agreement with single-cell RNA-sequencing data obtained from C. elegans which show near ubiquitous expression of flwr-1/F20D1.1 (Figure 1 – Figure supplement 2C, Table S3; Taylor et al. (2021)). The GFP::FLWR-1 fusion protein was functional as demonstrated by rescue of the swimming phenotype (recovery from cholinergic neuron photostimulation) observed in the flwr-1 mutant background (Figure 1 – Figure supplement 2D, E).

FLWR-1 is expressed in neurons and localizes to synaptic vesicles and the plasma membrane

Earlier version: (A-D) Confocal micrographs (maximum projection of z-stacks or single plane) of animals co-expressing pflwr-1::GFP::FLWR-1 and psnb-1::mCherry::SNB-1. (A) Overview of GFP::FLWR-1 expression. Arrows indicate dorsal and ventral nerve cords (DNC and VNC, respectively). Scale bar, 100 µm. (B) GFP::FLWR-1 in head neurons and pharynx. Scale bar, 20 µm. (C) Animal depicted in (B), single plane showing neck muscle cells. Arrows indicate GFP::FLWR-1 localization to the plasma membrane. Scale bar, 20 µm. (D) GFP and mCherry fluorescence in the DNC. Scale bar, 10 µm. (E) Line scan analysis of colocalization of GFP::FLWR-1 and mCherry::SNB-1 along the DNC as represented in (D). R2 as determined by Pearson correlation. a.u. = arbitrary units of fluorescence intensity. (F) Micrograph depicting GFP::FLWR-1 and mCherry::SNB-1 fluorescence in sublateral nerve cords and commissures. Arrowheads indicate synaptic puncta. Arrow points towards SV precursor travelling along commissure as shown in Supplementary Movie 1. Scale bar, 10 µm. (G) Kymograph representing the SV precursor indicated in (F) travelling along commissures. Scale bar, 2 µm. (H) Comparison of the ratio of DNC to VNC fluorescence of GFP::FLWR-1 and mCherry::SNB-1 in wild type and unc-104(e1265) mutant background. Mean (± SEM). Each dot represents a single animal. Two-way ANOVA with Šídák’s correction. *** p < 0.001. Number of animals imaged in N = 3 biological replicates: wild type = 33, unc-104 = 29.

Revised version: (A - D) Confocal micrographs (maximum projection of z-stacks or single plane) of animals co-expressing flwr-1p::GFP::FLWR-1 and psnb-1::mCherry::SNB-1. (A) Overview of GFP::FLWR-1 expression. Arrows indicate dorsal and ventral nerve cords (DNC and VNC, respectively). Scale bar, 100 µm. (B) GFP::FLWR-1 in head neurons and pharynx. Scale bar, 20 µm. (C) Animal depicted in (B), single plane showing neck muscle cells. Arrows indicate GFP::FLWR-1 localization to the plasma membrane. Scale bar, 20 µm. (D) GFP and mCherry fluorescence in the DNC. Scale bar, 10 µm. (E) Line scan analysis of colocalization of GFP::FLWR-1 and mCherry::SNB-1 along the DNC as represented in (D). R2 as determined by Pearson correlation. a.u. = arbitrary units of fluorescence intensity. (F) Micrograph depicting GFP::FLWR-1 and mCherry::SNB-1 fluorescence in sublateral nerve cords and commissures. Arrowheads indicate synaptic puncta. Arrow points towards SV precursor travelling along commissure as shown in Supplementary Movie 1. Scale bar, 10 µm. (G) Kymograph representing the SV precursor indicated in (F) travelling along commissures. Scale bar, 2 µm. (H) Comparison of the ratio of DNC to VNC fluorescence of GFP::FLWR-1 and mCherry::SNB-1 in wild type and unc-104(e1265) mutant background. Mean (± SEM). Each dot represents a single animal. Two-way ANOVA with Šídák’s correction. *** p < 0.001. Number of animals imaged in N = 3 biological replicates: wild type = 33, unc-104 = 29.

Within muscle cells, FLWR-1 was primarily targeted to the PM (Fig. 2C). As its D. melanogaster homolog was localized to synaptic vesicles (SVs) (Yao et al., 2009), we wondered whether FLWR-1 would co-localize with known SV markers such as SNB-1 synaptobrevin-1 (Calahorro and Izquierdo, 2018; Nonet, 1999). Indeed, GFP::FLWR-1 fluorescence largely overlapped with co-expressed mCherry::SNB-1 in the dorsal nerve cord (DNC; Fig. 2D, E). FLWR-1 further seemed to be enriched in fluorescent puncta in DNC and sublateral nerve cords indicating synaptic localization (Fig. 2D - F), however it was not restricted to synaptic regions only, meaning it is likely present also in the PM. Furthermore, we observed moving particles, probably SV precursors, which contained SNB-1 and FLWR-1, travelling along commissures between ventral nerve cord (VNC) and DNC (Fig. 2F, G and Movie S1). Anterograde transport of these precursors towards synapses depends on kinesin-3/UNC-104 (Hall and Hedgecock, 1991; Klopfenstein and Vale, 2004). We used a reduction-of-function allele (e1265) affecting interaction of UNC-104 with its cargo to investigate whether FLWR-1 is actively transported towards synapses (Cong et al., 2021). Indeed, animals lacking functional UNC-104 showed a reduced amount of axonal GFP fluorescence in the nerve ring and DNC while cell bodies were clearly visible (Figure 1 – Figure supplement 2F). The ratio of DNC to VNC fluorescence was significantly decreased in unc-104 mutants suggesting defective anterograde transport (Fig. 2H). Distribution of SNB-1 was similarly affected (Cuentas-Condori et al., 2023; Gally and Bessereau, 2003). Together, these results argue that FLWR-1 is expressed in neurons (as well as in muscles and other cell types) and is actively transported towards synapses.

GABAergic signaling is increased in flwr-1 knockout mutants

Earlier version: Since FLWR-1 is also expressed in body wall muscles (BWMs), we wondered whether the observed aldicarb resistance originates from reduced ACh reception by ACh receptors (AChRs) (Mahoney et al., 2006). To test this, we exposed animals lacking FLWR-1 to levamisole, an AChR agonist which induces paralysis by hyperexcitation and body contraction, similar to aldicarb (Gottschalk et al., 2005; Sattelle et al., 2002). We found no significant differences in the rate of paralysis between wild type and mutant animals (Figure 1 – Figure supplement 3A). However, transgenic expression of FLWR-1 in BWMs partially rescued the aldicarb resistance (Figure 1 – Figure supplement 3B). Therefore, loss of FLWR-1 might decrease the ACh response of muscle cells, which may contribute to the aldicarb resistance of flwr-1 mutants. However as muscles can also affect motor neurons by inhibitory retrograde signaling (Hu et al., 2012; Simon et al., 2008; Tong et al., 2017), we wanted to assess if the loss of FLWR-1 in muscle may also have effects on motor neurons. Thus, we rescued FLWR-1 in muscle and either cholinergic or GABergic neurons (Figure 1 – Figure supplement 3B). This did not show any additive effects to the pure neuronal rescues, thus FLWR-1 effects on muscle cell responses to cholinergic agonists must be cell-autonomous. Since muscles are activated by ACh via two different AChRs, the levamisole receptor and the nicotine-sensitive receptor (N-AChR), a homopentamer of ACR-16 subunits (Almedom et al. 2009; Richmond and Jorgensen, 1999), we addressed the possibility that FLWR-1 may regulate expression or function of N-AChRs in muscle, to affect phenotypes of aldicarb resistance. We performed aldicarb assay in the absence of the ACR-16 (Figure 1 – Figure supplement 3C), which showed that the two mutations, flwr-1(ok3128) and acr-16(ok789), had additive effects. Thus, FLWR-1 does not affect aldicarb resistance through downregulation of nAChRs, as otherwise the double mutant would not be more resistant than the flwr-1 single mutant.

Revised version: Since FLWR-1 is also expressed in body wall muscles (BWMs), we wondered whether the observed aldicarb resistance originates from reduced ACh reception by ACh receptors (AChRs) (Mahoney et al., 2006). To test this, we exposed animals lacking FLWR-1 to levamisole, an AChR agonist which induces paralysis by hyperexcitation and body contraction, similar to aldicarb (Gottschalk et al., 2005; Sattelle et al., 2002). We found no significant differences in the rate of paralysis between wild type and mutant animals (Figure 1 – Figure supplement 3A). However, transgenic expression of FLWR-1 in BWMs partially rescued the aldicarb resistance (Figure 1 – Figure supplement 3B). Therefore, loss of FLWR-1 might decrease the ACh response of muscle cells, which may contribute to the aldicarb resistance of flwr-1 mutants. However as muscles can also affect motor neurons by inhibitory retrograde signaling (Hu et al., 2012; Simon et al., 2008; Tong et al., 2017), we wanted to assess if the loss of FLWR-1 in muscle may also have effects on motor neurons. Thus, we rescued FLWR-1 in muscle and either cholinergic or GABergic neurons (Figure 1 – Figure supplement 3B). The combined rescues did not show any additive effects to the pure neuronal rescues, thus FLWR-1 effects on muscle cell responses to cholinergic agonists might be cell-autonomous. Yet, as cholinergic rescue alone did not overcompensate flwr-1 mutants’ resistance in aldicarb assays (Figure 3D), this interpretation is complicated. Since muscles are activated by ACh via two different AChRs, the levamisole receptor and the nicotine-sensitive receptor (N-AChR), a homopentamer of ACR-16 subunits (Almedom et al. 2009; Richmond and Jorgensen, 1999), we addressed the possibility that FLWR-1 may regulate expression or function of N-AChRs in muscle, to affect phenotypes of aldicarb resistance. We performed aldicarb assay in the absence of the ACR-16 (Figure 1 – Figure supplement 3C), which showed that the two mutations, flwr-1(ok3128) and acr-16(ok789), had additive effects. Thus, FLWR-1 does not affect aldicarb resistance through downregulation of nAChRs, as otherwise the double mutant would not be more resistant than the flwr-1 single mutant.

Added in the revised version: GABAergic signaling is increased in flwr-1 knockout mutants

Added in the revised version: (A) Representative confocal micrographs of an animal co-expressing flwr-1p::GFP::FLWR-1 and TagRFP::ELKS-1 in cholinergic motor neurons (unc-17(short) promoter). The region in the DNC used to acquire images shown in B and C is indicated; this position is posterior to the vulva, anterior is left. Scale bar, 10 µm. (B, C) DNCs in animals co-expressing flwr-1p::GFP::FLWR-1 and TagRFP::ELKS-1 in cholinergic motor neurons (unc-17(short) promoter), or in GABAergic neurons (unc-47 promoter; C), respectively. Scale bar, 5 µm. Examples of fluorescent puncta which contain either both, FLWR-1 and ELKS-1 (blue arrowheads), or FLWR-1 only (red arrowheads) are indicated. The same puncta are indicated in the respective analysis of signal density along the DNC in Figure 3 – Figure supplement 1A, B. (D) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb with cholinergic (unc-17p) and GABAergic (unc-47p) expression of FLWR-1 in flwr-1(ok3128) mutant background. N = 3 - 8 biological replicates. (E) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb with unc-47(e307) and unc-47(e307)/flwr-1(ok3128) double mutants. N = 3 biological replicates. Two-way ANOVA with Tukey’s correction in D, E. ns not significant, *** p < 0.001.

In addition to muscle expression, we also observed enrichment of FLWR-1 in fluorescent puncta indicating presynaptic localization (Fig. 3A, B). Moreover, FLWR-1 partially co-localized with the dense projection marker ELKS-1 in cholinergic, as well as in GABAergic neurons (Fig. 3A-C; Figure 3 – Figure supplement 1A-C; Dai et al. (2006); Kittelmann et al. (2013a)). Unlike ELKS-1, FLWR-1 could also be found in intersynaptic regions, though to a lesser extent than in synapses. This indicates that FLWR-1 might be involved in neurotransmission at both cholinergic and GABAergic NMJs, yet is not exclusively localized to active zones. The expression of FLWR-1 relative to ELKS-1 in either neuron type was not obviously biased to GABAergic or cholinergic neurons (Figure 3 – Figure supplement 1A-C). Thus, to determine whether the site of action of FLWR-1 in NMJ signaling is also evenly located to cholinergic and GABAergic neurons, we rescued FLWR-1 in each cell type of flwr-1 mutants, using the promoters of the vesicular transporters of ACh (UNC-17) or GABA (UNC-47), respectively. Surprisingly, expression of FLWR-1 in GABAergic, but not in cholinergic neurons, rescued aldicarb resistance (Fig. 3D). This was unexpected, since mutations affecting SV recycling commonly lead to a slowed replenishment of SVs and thus reduced ACh release (Salcini et al., 2001; Schuske et al., 2003; Yu et al., 2018). However, our results indicate that loss of FLWR-1 leads to increased release of GABA, which counteracts the aldicarb induced paralysis (Camara et al., 2019). Additional expression of FLWR-1 in BWMs did not change these results, suggesting that its role in neurons is more crucial in affecting aldicarb sensitivity (Figure 1 – Figure supplement 3B). To confirm this, we crossed flwr-1 mutants to animals lacking the vesicular GABA transporter UNC-47. This abolishes GABA release and induces hypersensitivity to aldicarb (Vashlishan et al., 2008). We observed that the additional mutation of unc-47 led to a complete loss of aldicarb resistance in flwr-1 mutants (Fig. 3E). No difference between unc-47 and flwr-1; unc-47 double mutants could be observed. However, as we used a high concentration of aldicarb, possible additional effects of the double mutant may have been masked. Thus, we also tested these animals at a lower concentration of aldicarb. unc-47; flwr-1 double mutants were hypersensitive to aldicarb compared to unc-47 single mutants, suggesting that ACh release is also increased by the loss of FLWR-1 (Figure 1 – Figure supplement 3D). In the absence of the compensatory increase in GABAergic signaling, this conveys aldicarb hypersensitivity. Jointly, these findings support the notion that it is primarily increased GABA signaling which is the main cause of aldicarb resistance in flwr-1 mutants, yet neurotransmission in motor neurons may be generally upregulated.

Removed in the revised version: GABAergic signaling is increased in flwr-1 knockout mutants.

Removed in the revised version: (A) Representative confocal micrographs of an animal co-expressing pflwr-1::GFP::FLWR-1 and TagRFP::ELKS-1 in cholinergic motor neurons (unc-17(short) promoter). The region in the DNC used to acquire images shown in B and C is indicated; this position is posterior to the vulva, anterior is left. Scale bar, 10 µm. (B, C) DNCs in animals co-expressing pflwr-1::GFP::FLWR-1 and TagRFP::ELKS-1 in cholinergic motor neurons (unc-17(short) promoter), or in GABAergic neurons (unc-47 promoter; C), respectively. Scale bar, 5 µm. Examples of fluorescent puncta which contain either both, FLWR-1 and ELKS-1 (blue arrowheads), or FLWR-1 only (red arrowheads) are indicated. The same puncta are indicated in the respective analysis of signal density along the DNC in Figure 3 – Figure supplement 1A, B. (D) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb with cholinergic (punc-17) and GABAergic (punc-47) expression of FLWR-1 in flwr-1(ok3128) mutant background. N = 3 - 8 biological replicates. (E) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb with unc-47(e307) and unc-47(e307)/flwr-1(ok3128) double mutants. N = 3 biological replicates. Two-way ANOVA with Tukey’s correction in D, E. ns not significant, *** p < 0.001.

Removed in the revised version: FLWR-1 localizes to cholinergic and GABAergic active zones.

Removed in the revised version: (A) Line scan analysis of colocalization of GFP::FLWR-1 and TagRFP::ELKS-1 along the DNC as represented in the same order in Fig. 3B. R2 as determined by Pearson correlation. a.u. = arbitrary units of fluorescence intensity. (B) Line scan analysis of colocalization of GFP::FLWR-1 and TagRFP::ELKS-1 along the DNC as represented in the micrographs in the same order in Fig. 3C. R2 as determined by Pearson correlation. a.u. = arbitrary units of fluorescence intensity. (A + B) Examples for fluorescent puncta which contain either both, FLWR-1 and ELKS-1 (blue arrowheads), or only FLWR-1 (red arrowheads) are indicated. The same puncta are indicated in the respective micrographs in Fig. 3B + C. (C) Comparison of Pearson correlation coefficients of line scans along DNCs represented in Fig. 3. Mean (± SEM). Unpaired t-test.

Depolarization-induced neuronal Ca2+ is increased in flwr-1 mutants

Earlier version: Since the release of GABA appeared to be increased in flwr-1 mutants, we wondered whether the response of GABAergic neurons during activation was increased. To test this, the fluorescent Ca2+ indicator GCaMP was expressed in GABAergic neurons, allowing us to estimate relative Ca2+ levels at presynaptic sites (Lu et al., 2022; Nakai et al., 2001). To allow optogenetic depolarization of neurons independent of GCaMP excitation light, we co-expressed the red-shifted channelrhodopsin variant ChrimsonSA (Oda et al., 2018; Seidenthal et al., 2022). As expected, optogenetic stimulation caused an increase in GCaMP fluorescence at NMJs in the DNC (Fig. 4A). Comparing wild type and flwr-1 mutants, we found that the gain in fluorescence intensity was significantly higher in animals lacking FLWR-1 (Fig. 4A, B), supporting our previous finding of increased neurotransmission in GABAergic neurons. To verify this at the behavioral level, we used a strain expressing ChR2(H134R) in GABAergic neurons, as it can be used to assess GABA release through measurement of body length (Liewald et al., 2008). Since GABA receptors hyperpolarize muscle cells, optogenetic stimulation of GABAergic motor neurons causes relaxation and thus increased body length (Schultheis et al., 2011; Seidenthal et al., 2022). To augment the effect of ChR2 stimulation on body length and to observe the effect of GABA release independent of cholinergic neurotransmission, the assay was performed in a mutant background lacking the levamisole receptor (unc-29(e1072) mutants, affecting an essential subunit; Fleming et al. (1997); Richmond and Jorgensen (1999)). As expected, stimulation with blue light led to increased body length (Fig. 4C). In accordance with Ca2+ imaging results, flwr-1; unc-29 double mutants showed a significantly stronger elongation during stimulation than unc-29 mutants (Fig. 4C, D). These results indicate that Ca2+ influx into the synapse is increased, likely causing more GABA to be released in animals lacking FLWR-1. The previously observed aldicarb resistance indicates that this change in GABAergic transmission outweighs possible changes in cholinergic transmission (Fig. 3).

Revised version: Since the release of GABA appeared to be increased in flwr-1 mutants, we wondered whether the response of GABAergic neurons during activation was increased. To test this, the fluorescent Ca2+ indicator GCaMP was expressed in GABAergic neurons, allowing us to estimate relative Ca2+ levels at presynaptic sites (Lu et al., 2022; Nakai et al., 2001). To allow optogenetic depolarization of neurons independent of GCaMP excitation light, we co-expressed the red-shifted channelrhodopsin variant ChrimsonSA (Oda et al., 2018; Seidenthal et al., 2022). As expected, optogenetic stimulation caused an increase in GCaMP fluorescence at NMJs in the DNC (Fig. 4A). Comparing wild type and flwr-1 mutants, we found that the gain in fluorescence intensity was significantly higher in animals lacking FLWR-1 (Fig. 4A, B), supporting our finding of increased neurotransmission in GABAergic neurons. To verify this at the behavioral level, we used a strain expressing ChR2(H134R) in GABAergic neurons, as it can be used to assess GABA release through measurement of body length (Liewald et al., 2008). Since GABA receptors hyperpolarize muscle cells, optogenetic stimulation of GABAergic motor neurons causes relaxation and thus increased body length (Schultheis et al., 2011; Seidenthal et al., 2022). To augment the effect of ChR2 stimulation on body length and to observe the effect of GABA release independent of cholinergic neurotransmission, the assay was performed in a mutant background lacking the levamisole receptor (unc-29(e1072) mutants, affecting an essential subunit; Fleming et al. (1997); Richmond and Jorgensen (1999)). As expected, stimulation with blue light led to increased body length (Fig. 4C). In accordance with Ca2+ imaging results, flwr-1; unc-29 double mutants showed a significantly stronger elongation during stimulation than unc-29 mutants (Fig. 4C, D). These results indicate that Ca2+ influx into the synaptic cytosol is increased, likely causing more GABA to be released in animals lacking FLWR-1. The previously observed aldicarb resistance indicates that this change in GABAergic transmission outweighs possible changes in cholinergic transmission (Fig. 3).

Loss of FLWR-1 leads to increased Ca2+ levels during stimulation

(A) Mean (± SEM) normalized fluorescence in synaptic puncta of the DNC of animals expressing GCaMP3 and ChrimsonSA in GABAergic neurons (unc-25 promoter). All animals were supplemented with ATR. A 10 s light pulse (590 nm, 40 µW/mm2) was applied after 5 s as indicated by the red shade. (B) Mean (± SEM) normalized fluorescence during stimulation (seconds 6 to 14) as depicted in (A). Each dot indicates a single animal. Unpaired t-test. ** p < 0.01. (A + B) Number of animals imaged in N = 5 biological replicates: wild type = 40, flwr-1 = 41. Outliers were removed from both datasets as detected by the iterative Grubb’s method (GraphPad Prism). (C) Mean (± SEM) body length of animals expressing ChR2(H134R) in GABAergic neurons (unc-47 promoter) in the unc-29(e1072) mutant background, normalized to the average before stimulation. All animals were supplemented with ATR. A 20 s light pulse (470 nm, 100 µW/mm2) was applied after 5 s as indicated by the blue shade. (D) Mean (± SEM) relative body length during stimulation (seconds 6 to 24) as depicted in (C). Each dot indicates a single animal. Unpaired t-test. * p < 0.05. Number of animals measured in N = 4 biological replicates: wild type = 51, flwr-1 = 49. (E) Mean (± SEM) normalized fluorescence in synaptic puncta of the DNC of animals expressing GCaMP6f and ChrimsonSA in cholinergic motor neurons (unc-17b promoter). All animals were supplemented with ATR. A 10 s light pulse (590 nm, 40 µW/mm2) was applied after 5 s as indicated by red shade. (F) Median (with IQR) normalized fluorescence during stimulation (seconds 6 to 14) as depicted in (E). Each dot indicates a single animal. Kruskal-Wallis test. Only statistically significant differences are depicted. * p < 0.05, ** p < 0.01. (E + F) Number of animals imaged in N = 5 biological replicates: wild type = 40, flwr-1 = 38, cholinergic rescue = 39, muscle rescue = 36. No outliers were detected by the iterative Grubb’s method. (G) Schematic representation of motor neuron innervation of BWMs. Arrows indicate the putatively increased (green) or decreased (red) neurotransmission/excitation of the involved cell types in flwr-1 mutants compared to wildtype.

Removed in the revised version: Basal SNG-1::pHluorin fluorescence and cell surface fraction are unchanged in flwr-1 mutants.

Removed in the revised version: (A) Mean (± SEM) pHluorin fluorescence before stimulation as depicted in Fig. 5. Number of animals imaged in N = 5 biological replicates: wild type = 27, flwr-1 = 32. Unpaired t-test. ns not significant p > 0.05. (B) Representative images of a primary neuronal cell expressing SNG-1::pHluorin when exposed to different buffers as indicated. Scale bar, 5 µm. (C) Number of cells imaged in N = 2 biological replicates: wild type = 29, flwr-1 = 23. Unpaired t-test. ns not significant p > 0.05.

Earlier version: To investigate whether neuronal responses to depolarization may be generally upregulated in flwr-1 mutants, we assessed whether evoked Ca2+ influx is affected in cholinergic motor neurons as well (Fig. 4E, F). Again, we observed increased GCaMP fluorescence levels during optogenetic stimulation. This effect could be rescued in cholinergic neurons. However, since postsynaptic muscle exerts inhibitory retrograde signaling to presynaptic cholinergic neurons (Hu et al., 2012; Simon et al., 2008; Tong et al., 2017), and because FLWR-1 is also expressed in muscle, we tested if the phenotype of the loss of FLWR-1 in cholinergic neurons would be affected by rescuing FLWR-1 in muscle. This was not the case (Fig. 4E, F). Together, these results indicate that loss of FLWR-1 conveys an upregulation of neuronal Ca2+ influx during continuous stimulation in both classes of motor neurons (Fig. 4G). However, the overall E/I balance appears to be shifted towards stronger inhibition of BWMs.

Revised version: To investigate whether neuronal responses to depolarization may be generally upregulated in flwr-1 mutants, we assessed whether evoked Ca2+ level increase is affected in cholinergic motor neurons as well (Fig. 4E, F). Again, we observed increased GCaMP fluorescence levels during optogenetic stimulation. This effect could be rescued in cholinergic neurons. However, since postsynaptic muscle exerts inhibitory retrograde signaling to presynaptic cholinergic neurons (Hu et al., 2012; Simon et al., 2008; Tong et al., 2017), and because FLWR-1 is also expressed in muscle, we tested if the phenotype of the loss of FLWR-1 in cholinergic neurons would be affected by rescuing FLWR-1 in muscle. This was not the case (Fig. 4E, F). Together, these results indicate that loss of FLWR-1 conveys an upregulation of neuronal Ca2+ level rise during continuous stimulation in both classes of motor neurons (Fig. 4G). However, the overall E/I balance appears to be shifted towards stronger inhibition of BWMs.

Endocytosis is slowed in flwr-1 mutants in non-neuronal cells and neurons

Earlier version: Homologues of FLWR-1 were implicated in endocytosis in neurons as well as in non-neuronal cells (Chang et al., 2018; Rudd et al., 2023; Yao et al., 2009). We thus wondered whether this function is evolutionarily conserved in nematodes. One possibility to assess this in C. elegans is to observe endocytosis in coelomocytes (CCs; (Fares and Greenwald, 2001). These scavenger cells continuously endocytose fluid from the body cavity and loss of endocytosis-associated factors affects uptake of proteins secreted from other tissues (Fares and Grant, 2002). GFP fused to a secretory signal sequence is discharged from BWMs, and its endocytic uptake by CCs can be quantified by fluorescence microscopy (Bednarek et al., 2007; Fares and Greenwald, 2001). According to single-cell RNAseq data, FLWR-1 is expressed in CCs (Table S3; Taylor et al. (2021). Indeed, mutation of flwr-1 led to strongly reduced GFP fluorescence levels within CCs indicating a reduced uptake by endocytosis (Fig. 5A, B). This defect could be cell-autonomously rescued by expressing FLWR-1 in CCs from the unc-122 promoter. Furthermore, to assess recycling of SVs, we used the pOpsicle (pH-sensitive optogenetic reporter of synaptic vesicle recycling) assay we recently established (Seidenthal et al., 2023) to estimate the amount of SV fusion and the rate of recycling of SV components in cholinergic neurons. This assay combines a pHluorin based probe fused to an SV-associated protein (SNG-1) and optogenetic stimulation of neurotransmitter release; this way, pHluorin fluorescence is unquenched during stimulated exocytosis and quenched during SV endocytosis and recycling (Sankaranarayanan et al., 2000). Since in this assay signals originating from SVs and from the PM contribute to the overall signal, and since FLWR-1 has an effect on endocytosis, we needed to verify that the relative localization of the SNG-1::pHluorin sensor itself was not affected by the flwr-1 mutation. Basal fluorescence, before stimulation, was unaltered in flwr-1 mutants, showing that there is no FLWR-1 dependent alteration of SNG-1::pHluorin in cellular membranes (Figure 4). In addition, to estimate the amount of the sensor present in the PM vs. in vesicles, we used primary culture of C. elegans cells. Cholinergic neurons expressing SNG-1::pHluorin were imaged and exposed to different pH by adding a buffer of pH 5.6, or by adding ammonium chloride buffer of physiological / neutral pH, which can penetrate the cell and thus shows the maximum achievable signal, i.e. all unquenched pHluorin signal present in the cell. This showed that the amount of SNG-1::pHluorin present in SVs was also not affected by the flwr-1 mutation (Figure 4 – Figure supplement 1B, C). In vivo, loss of FLWR-1 led to significantly increased fluorescence signals during stimulation which indicates more SV fusion compared to wild type (Fig. 5C, E). This is in accordance with the increased responses of cholinergic neurons to depolarization we observed earlier (Fig. 4E, F). Moreover, the rate of fluorescence decay after stimulation was significantly reduced in flwr-1 mutants suggesting slower recycling of SVs, or reduced acidification of endosomal structures or SVs after recycling (Fig. 5D, F).

Revised version: Homologues of FLWR-1 were implicated in endocytosis in neurons as well as in non-neuronal cells (Chang et al., 2018; Rudd et al., 2023; Yao et al., 2009). We thus wondered whether this function is evolutionarily conserved in nematodes. One possibility to assess this in C. elegans is to observe endocytosis in coelomocytes (CCs; (Fares and Greenwald, 2001). These scavenger cells continuously endocytose fluid from the body cavity and loss of endocytosis-associated factors affects uptake of proteins secreted from other tissues (Fares and Grant, 2002). GFP fused to a secretory signal sequence is discharged from BWMs, and its endocytic uptake by CCs can be quantified by fluorescence microscopy (Bednarek et al., 2007; Fares and Greenwald, 2001). According to single-cell RNAseq data, FLWR-1 is expressed in CCs (Table S3; Taylor et al. (2021). Indeed, mutation of flwr-1 led to strongly reduced GFP fluorescence levels within CCs indicating a reduced uptake by endocytosis (Fig. 5A, B). This defect could be cell-autonomously rescued by expressing FLWR-1 in CCs from the unc-122 promoter. Furthermore, to assess recycling of SVs, we used the pOpsicle (pH-sensitive optogenetic reporter of synaptic vesicle recycling) assay we recently established (Seidenthal et al., 2023) to estimate the amount of SV fusion and the rate of recycling of SV components in cholinergic neurons. This assay combines a pHluorin based probe fused to an SV-associated protein (SNG-1) and optogenetic stimulation of neurotransmitter release; this way, pHluorin fluorescence is unquenched during stimulated exocytosis and quenched during SV endocytosis and recycling (Sankaranarayanan et al., 2000). Since in this assay signals originating from SVs and from the PM contribute to the overall signal, and since FLWR-1 has an effect on endocytosis, we needed to verify that the relative localization of the SNG-1::pHluorin sensor itself was not affected by the flwr-1 mutation. Basal fluorescence, before stimulation, was unaltered in flwr-1 mutants, showing that there is no FLWR-1 dependent alteration of SNG-1::pHluorin in cellular membranes (Figure 4 – Figure supplement 1A). In addition, to estimate the amount of the sensor present in the PM vs. in vesicles, we used primary culture of C. elegans cells. Cholinergic neurons expressing SNG-1::pHluorin were imaged and exposed to different pH by adding a buffer of pH 5.6, or by adding ammonium chloride buffer of physiological / neutral pH, which can penetrate the cell and thus shows the maximum achievable signal, i.e. all unquenched pHluorin signal present in the cell. This showed that the amount of SNG-1::pHluorin present in SVs was also not affected by the flwr-1 mutation (Figure 4 – Figure supplement 1B, C). In vivo, loss of FLWR-1 led to significantly increased fluorescence signals during stimulation which indicates more SV fusion compared to wild type (Fig. 5C, E). This is in accordance with the increased responses of cholinergic neurons to depolarization we observed earlier (Fig. 4E, F). Moreover, the rate of fluorescence decay after stimulation was significantly reduced in flwr-1 mutants suggesting slower recycling of SVs, or reduced acidification of endosomal structures or SVs after recycling (Fig. 5D, F).

FLWR-1 facilitates endocytosis in non-neuronal and neuronal cells

(A) Exemplary images of the coelomocytes (CCs) in wild type, flwr-1(ok3128) and in flwr-1 mutants expressing FLWR-1 in CCs (unc-122 promoter). GFP containing a secretion signal sequence (ssGFP) was expressed in BWMs (myo-3 promoter). Scale bar, 10 µm. (B) Median (with IQR) normalized fluorescence of CCs. Each dot indicates a single CC. Kruskal-Wallis test. ** p < 0.001*** p < 0.001. Number of CCs imaged in N = 3 biological replicates: wild type = 158, flwr-1 = 185, rescue = 113. (C) Mean (± SEM) normalized DNC fluorescence of animals expressing SNG-1::pHluorin and ChrimsonSA in cholinergic neurons (unc-17 promoter). All animals were supplemented with ATR. A 30 s light pulse (590 nm, 40 µW/mm2) was applied after 10 s as indicated by the red shade. (D) Mean (± SEM) pHluorin fluorescence as depicted in (C) but additionally normalized to the maximum value of each dataset. (E) Mean (± SEM) normalized fluorescence during stimulation (seconds 15 to 35) as depicted in (C). Each dot indicates a single animal. Unpaired t-test. (F) Mean (± SEM) calculated exponential decay constants of fluorescence decline after stimulation. Each dot indicates a single animal. Unpaired t-test. (C – F) Number of animals imaged in N = 5 biological replicates: wild type = 27, flwr-1 = 32. * p < 0.05. (G) Representative voltage-clamp recordings of currents in BWMs. Animals expressing ChR2(H134R) in cholinergic motor neurons (unc-17 promoter, transgene zxIs6) were treated with ATR. A 30 s light stimulus (470 nm, 8 mW/mm2) was applied as indicated by blue bars. (H) Normalized mPSC frequency in BWMs. All animals were treated with ATR. A 30 s light pulse (470 nm, 8 mW/mm2) was applied as indicated by the blue shade. Dashed lines indicate one-phase exponential regression analysis fitted to the mean mPSC frequencies during stimulation. Calculated time constants of decay are shown. Two-way ANOVA with Šidák’s correction. All significant differences to wild type are depicted. (I) mPSC amplitude in BWMs of animals measured in (G + H). (J) Mean (± SEM) mPSC amplitude during light stimulation as indicated in (I). Unpaired t-test. * p < 0.05. (G - J) Number of animals: wild type = 9, flwr-1 = 8.

Removed in the revised version: flwr-1 mutants show defecting cholinergic neurotransmission only during continuous stimulation.

Removed in the revised version: (A,B) Mean (± SEM) mPSC frequency and amplitude, respectively, before stimulation. Unpaired t-test. ns not significant. Number of animals: wild type = 16, flwr-1 = 14. (C) Mean (± SEM) inward currents of BWMs recordings induced by 10 ms light pulses (470 nm, 8 mW/mm2) applied every 2 s (0.5 Hz). Two-way ANOVA with Šidák’s correction for multiple comparisons. ns not significant. Number of animals: wild type = 9, flwr-1 = 8. (D) As in (C), but 2Hz stimulation. ns not significant. Number of animals: wild type = 7, flwr-1 = 7. (E) Representative voltage-clamp recording of currents detected in BWMs. This wild type animal expresses ChR2(H134R) in cholinergic motor neurons (unc-17 promoter) and has been treated with ATR. A 30 s light stimulus (470 nm, 8 mW/mm2) as well as a 10 ms pulse after a 15 s inter-stimulus interval (ISI) were applied as indicated by blue bars. (F) Mean (± SEM) mPSC frequency in BWMs of animals expressing ChR2(H134R) in cholinergic motor neurons (unc-17 promoter). All animals have been treated with ATR. A 30 s light pulse (470 nm, 8 mW/mm2) was applied as indicated by blue shade. (G) Mean (± SEM) mPSC frequency during 30 s stimulation. Unpaired t-test. ** p < 0.01. (H) Analysis of the amplitude of the first peak during 30 s photostimulation and the second peak after 15 s ISI as indicated in (E). Two-way ANOVA with Tukey’s correction for multiple comparisons. ns not significant, *** p < 0.001. (F - H) Number of animals: wild type = 9, flwr-1 = 8.

Earlier version: Slowed replenishment of SV pools caused by defective recycling is known to cause synaptic depression, as assessed by electrophysiological measurements of post-synaptic currents (Kittelmann et al., 2013b; Krick et al., 2021a; Wu and Betz, 1998). This is also the case in Drosophila for animals lacking Flower (Yao et al., 2009). To address this in C. elegans, we measured miniature post-synaptic currents (mPSCs; minis) in BWMs (Fig. 5G), which reflect the post-synaptic effects of presynaptic neurotransmitter release (Liewald et al., 2008; Weissenberger et al., 2011). Animals lacking FLWR-1 showed no significant difference in basal mPSC frequency or amplitude (Fig. 5H, I; Figure 5 – Figure supplement 1A, B) which is in accordance with FLWR-1 being dispensable in basal swimming locomotion (Fig. 1C, D). Similarly, pulsed optogenetic stimulation of cholinergic neurons at different frequencies did not reveal a difference in the measured currents between wild type and mutants (Figure 5 – Figure supplement 1C, D), suggesting that FLWR-1 might be needed only during continuous stimulation. Indeed, flwr-1 mutants showed an accelerated rundown of the mPSC frequency in BWMs during constant illumination (Fig. 5G, H). This is in accordance with a role of FLWR-1 in SV recycling. Surprisingly, mPSC amplitudes in flwr-1 mutants are reduced during 30 s hyperstimulation (Fig. 5I, J). This suggests that either the amount of neurotransmitter released from a single SV, or the number of multivesicular fusion events, is decreased (Liu et al., 2005). The absolute mPSC frequency, which represents the number of SVs fusing per time, was also smaller during continuous stimulation in flwr-1 animals (Figure 5 – Figure supplement 1E-G). These results contrast the increased Ca2+ responses of cholinergic neurons we observed (Fig. 4E, F). Since cholinergic neurons also stimulate GABAergic neurons, and since GABAergic minis are also evaluated here, the dissection of animals, required for electrophysiological recordings, could have damaged neuronal commissures, causing an interruption of physiological signal transmission, which is otherwise observed in intact animals. We further analyzed whether flwr-1 mutants can recover from strong optogenetic 30 s stimulation by applying a short stimulus after a recovery period (inter-stimulus interval, ISI) of 15 s, and found no significant difference to wild type (Figure 5 – Figure supplement 1E, H). This suggests that a 15 s ISI is sufficient for flwr-1 mutants to recover to the same extent as wild type synapses. In sum, our results support a facilitatory role of FLWR-1 in SV recycling specifically during continuous stimulation.

Revised version: Slowed replenishment of SV pools caused by defective recycling is known to cause synaptic depression, as assessed by electrophysiological measurements of post-synaptic currents (Kittelmann et al., 2013b; Krick et al., 2021; Wu and Betz, 1998). This is also the case in Drosophila for animals lacking Flower (Yao et al., 2009). To address this in C. elegans, we measured miniature post-synaptic currents (mPSCs; minis) in BWMs (Fig. 5G), which reflect the post-synaptic effects of presynaptic neurotransmitter release (Liewald et al., 2008; Weissenberger et al., 2011). Animals lacking FLWR-1 showed no significant difference in basal mPSC frequency or amplitude (Fig. 5H, I; Figure 5 – Figure supplement 1A, B) which is in accordance with FLWR-1 being dispensable in basal swimming locomotion (Fig. 1C, D). Similarly, pulsed optogenetic stimulation of cholinergic neurons at different frequencies did not reveal a difference in the measured currents between wild type and mutants (Figure 5 – Figure supplement 1C, D), suggesting that FLWR-1 might be needed only during continuous stimulation. Indeed, flwr-1 mutants showed an accelerated rundown of the mPSC frequency in BWMs during constant illumination (Fig. 5G, H). This is in accordance with a role of FLWR-1 in SV recycling. Surprisingly, mPSC amplitudes in flwr-1 mutants are reduced during 30 s hyperstimulation (Fig. 5I, J). This suggests that either the amount of neurotransmitter released from a single SV, or the number of multivesicular fusion events, is decreased (Liu et al., 2005). The absolute mPSC frequency, which represents the number of SVs fusing per time, was also smaller during continuous stimulation in flwr-1 animals (Figure 5 – Figure supplement 1E-G). These results contrast the increased Ca2+ responses of cholinergic neurons we observed (Fig. 4E, F). Since cholinergic neurons also stimulate GABAergic neurons, and since GABAergic minis are also evaluated here, the dissection of animals, required for electrophysiological recordings, could have damaged neuronal commissures, causing an interruption of physiological signal transmission, which is otherwise observed in intact animals. We further analyzed whether flwr-1 mutants can recover from strong optogenetic 30 s stimulation by applying a short stimulus after a recovery period (inter-stimulus interval, ISI) of 15 s, and found no significant difference to wild type (Figure 5 – Figure supplement 1E, H). This suggests that a 15 s ISI is sufficient for flwr-1 mutants to recover to the same extent as wild type synapses. In sum, our results support a facilitatory role of FLWR-1 in SV recycling specifically during continuous stimulation.

Loss of FLWR-1 leads to depleted SV pools and accumulation of endocytic structures post stimulation

Earlier version: Mutants in which endocytosis is affected commonly have fewer SVs because of defective recovery of SV components; this was also found for presynaptic boutons in Drosophila mutants lacking Flower (Schuske et al., 2003; Yao et al., 2009). Such a defect is even more pronounced when samples are conserved immediately following optogenetic stimulation by high-pressure freezing (Kittelmann et al., 2013b; Weimer, 2006; Yu et al., 2018). In nematodes, this optogenetic stimulation can be controlled by comparing animals supplemented with ATR, the ChR2 chromophore (Nagel et al., 2005), to animals without ATR. Ultrastructural analysis using transmission electron microscopy (TEM) indeed revealed fewer SVs in stimulated flwr-1 mutant synapses compared to wild type (Fig. 6A, B). This indicates that flwr-1 mutants, unlike wild type, are unable to refill SV pools sufficiently fast. At the same time, Drosophila Flower mutants were shown to be defective in formation of bulk endosomal structures after strong stimulation (Yao et al., 2017). In contrast to this, we observed an increased, rather than decreased, number of endocytic structures (large vesicles / ‘100 nm vesicles’) in stimulated flwr-1 mutant synapses (Fig. 6C). This may either be caused by increased SV fusion which triggers bulk endosomal formation (Clayton et al., 2008; Wu et al., 2014) or defective breakdown of these endocytic structures (Gan and Watanabe, 2018; Watanabe et al., 2013; Yu et al., 2018). Previously, we observed the formation of very large endocytic structures in mutants that affect their resolution into SVs, like endophilin or synaptojanin (Kittelmann et al., 2013b), while a dynamin mutant showed unresolved endocytic structures at, and in continuity with, the PM. Our pHluorin imaging data (Fig. 5C - F), would support the notion that both, increased SV fusion as well as defective recovery and subsequent acidification of SVs, may cause the higher number of endocytic structures in flwr-1 mutants. However, we note that our stimulation regime likely resembles a more physiological activation of neurotransmission compared to the intense stimuli previously used (Yao et al., 2017), as wild type synapses only rarely contained endocytic structures (Kittelmann et al., 2013b). Moreover, we observed fewer docked vesicles in flwr-1 animals which have been treated with ATR, compared to those without (Fig. 6D). This might be caused by increased SV fusion and is in line with the increased Ca2+-levels during stimulation we observed before. The number of neuropeptide-containing dense core vesicles (DCVs) was unchanged in animals lacking FLWR-1 (Figure 6 – Figure supplement 1).

Revised version: Mutants in which endocytosis is affected commonly have fewer SVs because of defective recovery of SV components; this was also found for presynaptic boutons in Drosophila mutants lacking Flower (Schuske et al., 2003; Yao et al., 2009). Such a defect is even more pronounced when samples are conserved immediately following optogenetic stimulation by high-pressure freezing (Kittelmann et al., 2013b; Weimer, 2006; Yu et al., 2018). In nematodes, this optogenetic stimulation can be controlled by comparing animals supplemented with ATR, the ChR2 chromophore (Nagel et al., 2005), to animals without ATR. Ultrastructural analysis using transmission electron microscopy (TEM) indeed revealed fewer SVs in stimulated flwr-1 mutant synapses compared to wild type (Fig. 6A, B). This indicates that flwr-1 mutants, unlike wild type, are unable to refill SV pools sufficiently fast. At the same time, Drosophila Flower mutants were shown to be defective in formation of bulk endosomal structures after strong stimulation (Yao et al., 2017). In contrast to this, we observed an increased, rather than decreased, number of endocytic structures (large vesicles / ‘100 nm vesicles’) in stimulated flwr-1 mutant synapses (Fig. 6C), that were of larger size (Figure 6 – Figure supplement 1B). This may either be caused by increased SV fusion which triggers bulk endosomal formation (Clayton et al., 2008; Wu et al., 2014) or defective breakdown of these endocytic structures (Gan and Watanabe, 2018; Watanabe et al., 2013; Yu et al., 2018). Previously, we observed the formation of very large endocytic structures in mutants that affect their resolution into SVs, like endophilin or synaptojanin (Kittelmann et al., 2013b), while a dynamin mutant showed unresolved endocytic structures at, and in continuity with, the PM. Our pHluorin imaging data (Fig. 5C - F), would support the notion that both, increased SV fusion as well as defective recovery and subsequent acidification of SVs, may cause the higher number of endocytic structures in flwr-1 mutants. However, we note that our stimulation regime likely resembles a more physiological activation of neurotransmission compared to the intense stimuli previously used (Yao et al., 2017), as wild type synapses only rarely contained endocytic structures (Kittelmann et al., 2013b). Moreover, we observed fewer docked vesicles in flwr-1 animals which have been treated with ATR, compared to those without (Fig. 6D). This might be caused by increased SV fusion and is in line with the increased Ca2+-levels during stimulation we observed before. The number of neuropeptide-containing dense core vesicles (DCVs) was unchanged in animals lacking FLWR-1 (Figure 6 – Figure supplement 1A).

Ultra-structural analysis reveals defective recycling of SVs after stimulation in flwr-1 mutants

(A) Representative TEM micrographs of cholinergic en-passant synapses in wild type and flwr-1(ok3128) animals expressing ChR2(H134R) in cholinergic neurons (unc-17 promoter). Animals were optionally treated with ATR as indicated. Dense projections (DP), endosomes (E), dense core vesicles (blue arrows), SVs (black arrowheads), docked vesicles (white arrowheads) and large vesicles (LV) are indicated. Scale bars, 100 nm. (B) Violin plot depicting the number of SVs counted per synaptic profile. (C) Violin plot depicting the number of large endcytic vesicels and ‘endosomes’ per synapse. (D) Violin plot depicting the number of docked vesicles observed per synaptic profile. (B - D) Bold line represents the median, and the dashed lines the IQR. Kruskal-Wallis test. Only statistically significant differences are depicted. * p < 0.05,** p < 0.01, *** p < 0.001. Number of synaptic profiles imaged: wild type (-ATR) = 56, wild type (-ATR) = 51, flwr-1 (-ATR) = 55, flwr-1 (+ATR) = 59.

Removed in the revised version: flwr-1 mutants have a normal numbers of dense core vesicles before and after stimulation.

Removed in the revised version: Violin plot depicting the number of dense core vesicles counted per synaptic profile. Bold line represents the median and the dashed lines the IQR. Kruskal-Wallis test. ns not significant. Number of synaptic profiles imaged: wild type (-ATR) = 56, wild type (-ATR) = 51, flwr-1 (-ATR) = 55, flwr-1 (+ATR) = 59.

The increased Ca2+-levels of flwr-1 neurons may be caused by deregulation of MCA-3

Earlier version: While a facilitating role of Flower in endocytosis appears to be conserved in C. elegans, we found no evidence that FLWR-1 conducts Ca2+ upon insertion into the PM (Yao et al., 2009). On the contrary, presynaptic Ca2+ levels were increased during photostimulation of neurons in animals lacking FLWR-1. We thus wondered whether clearance of Ca2+ which has entered the synapse via VGCCs, might be defective in flwr-1 mutants. In neurons and muscles, the plasma membrane Ca2+ ATPase (PMCA) is involved in extruding Ca2+ from the cell (Boczek et al., 2019; Krebs, 2022; Krick et al., 2021a). The C. elegans homolog MCA-3 (also known as CUP-7) is expressed in neurons, muscle cells and CCs (Figure 7 – Figure supplement 1A, Table S3; Bednarek et al. (2007); Taylor et al. (2021)). Interestingly, reducing the function of MCA-3 by mutation was shown to cause a similar defect in endocytosis of secreted GFP in CCs as the loss of FLWR-1 does (Fig. 5A, B; Bednarek et al. (2007)). We therefore wondered whether MCA-3 might be negatively affected in flwr-1 mutants. To assess this, we used a mutant mca-3(ok2048) lacking part of the C-terminal, regulatory calmodulin-binding domain (Consortium, 2012; Kraev et al., 1999; Mantilla et al., 2023) Fig. 7A). Since mca-3 loss-of-function (l.o.f.) mutants are lethal, it is likely that this represents a reduction-of-function (r.o.f.) mutation (Bednarek et al., 2007). As expected, mca-3 mutants displayed increased Ca2+ influx (or net Ca2+ levels, representing the summed VGCC-mediated entry and MCA-3-mediated efflux) upon optogenetic stimulation of cholinergic motor neurons, similar to the effect of the flwr-1 mutation (Fig. 7B, C). Elevated Ca2+ levels were not further enhanced in a flwr-1; mca-3 double mutant. Our data suggest that the two genes are acting in a common pathway. A partially different picture was observed in aldicarb assays, as reduction of MCA-3 function conveyed aldicarb resistance that was, however, less pronounced than for flwr-1 mutants (Fig. 7D). Nevertheless, since the double mutant showed no exacerbated phenotype compared to the flwr-1 mutant, both proteins appear to function in the same pathway. If MCA-3 function was augmented by FLWR-1, then the loss of FLWR-1 may be overcome by higher MCA-3 expression. Since the main focus of FLWR-1 was the GABAergic NMJ (Fig. 3D, E), we analyzed aldicarb resistance in flwr-1 mutants in which we overexpressed MCA-3 in GABAergic neurons (Figure 7 – Figure supplement 1B). Indeed, MCA-3 overexpression in GABAergic neurons efficiently rescued the flwr-1 resistance to wild type levels.

Revised version: While a facilitating role of Flower in endocytosis appears to be conserved in C. elegans, in contrast to previous findings from Drosophila (Yao et al., 2009), we found no evidence that FLWR-1 conducts Ca2+ upon insertion into the PM. On the contrary, presynaptic Ca2+ levels were increased during photostimulation of neurons in animals lacking FLWR-1. We thus wondered whether clearance of Ca2+ which has entered the synapse via VGCCs, might be defective in flwr-1 mutants. In neurons and muscles, the plasma membrane Ca2+ ATPase (PMCA) is involved in extruding Ca2+ from the cell (Boczek et al., 2019; Krebs, 2022; Krick et al., 2021). The C. elegans homolog MCA-3 (also known as CUP-7) is expressed in neurons, muscle cells and CCs (Figure 7 – Figure supplement 1A, Table S3; Bednarek et al. (2007); Taylor et al. (2021)). Interestingly, reducing the function of MCA-3 by mutation was shown to cause a similar defect in endocytosis of secreted GFP in CCs as the loss of FLWR-1 does (Fig. 5A, B; Bednarek et al. (2007)). We therefore wondered whether MCA-3 might be negatively affected in flwr-1 mutants. To assess this, we used a mutant mca-3(ok2048) lacking part of the C-terminal, regulatory calmodulin-binding domain (Consortium, 2012; Kraev et al., 1999; Mantilla et al., 2023) Fig. 7A). Since mca-3 loss-of-function (l.o.f.) mutants are lethal, it is likely that this represents a reduction-of-function (r.o.f.) mutation (Bednarek et al., 2007). As expected, mca-3 mutants displayed increased Ca2+ influx (or net Ca2+ levels, representing the summed VGCC-mediated entry and MCA-3-mediated efflux) upon optogenetic stimulation of cholinergic motor neurons, similar to the effect of the flwr-1 mutation (Fig 7.B, C). Elevated Ca2+ levels were not further enhanced in a flwr-1; mca-3 double mutant. Our data suggest that the two genes are acting in a common pathway. A partially different picture was observed in aldicarb assays, as reduction of MCA-3 function conveyed aldicarb resistance that was, however, less pronounced than for flwr-1 mutants (Fig. 7D). Nevertheless, since the double mutant showed no exacerbated phenotype compared to the flwr-1 mutant, both proteins appear to function in the same pathway. If MCA-3 function was augmented by FLWR-1, then the loss of FLWR-1 may be overcome by higher MCA-3 expression. Since the main focus of FLWR-1 was the GABAergic NMJ (Fig. 3D, E), we analyzed aldicarb resistance in flwr-1 mutants in which we overexpressed MCA-3 in GABAergic neurons (Figure 7 – Figure supplement 1B). Indeed, MCA-3 overexpression in GABAergic neurons efficiently rescued the flwr-1 resistance to wild type levels.

Increased Ca2+ levels in flwr-1 mutants may be caused by negative regulation of MCA-3

(A) Schematic representation of the mca-3 gene locus including exon/intron structure of isoforms mca-3a and mca-3b. Bars represent exons and connecting lines introns. The size of the ok2048 deletion as well as the putative calmodulin-binding domain are indicated. (B) Mean (± SEM) normalized fluorescence in synaptic puncta of the DNC of animals expressing GCaMP6f and ChrimsonSA in cholinergic motor neurons (unc-17b promoter). All animals were supplemented with ATR. A 10 s light pulse (590 nm, 40 µW/mm2) was applied after 5 s as indicated by the red shade. (C) Median (with IQR) normalized fluorescence during stimulation (seconds 6 to 14) as depicted in (B). Each dot indicates a single animal. Kruskal-Wallis test. Only statistically significant differences are depicted. ** p < 0.01, *** p < 0.001. Number of animals imaged in (B + C): wild type = 83, flwr-1 = 58, mca-3 = 50, mca-3; flwr-1 = 44. Outliers were removed from all datasets as detected by iterative Grubb’s method (GraphPad Prism). (D) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb. Two-way ANOVA with Tukey’s correction. ns not significant, *** p < 0.001.

Removed in the revised version: Comparison of flwr-1 and mca-3 expression by single-cell RNAseq data.

Removed in the revised version: (A) Heatmap plot depicting flwr-1(F20D1.1) and mca-3 single-cell RNAseq data generated by the CeNGEN database (https://cengen.shinyapps.io/CengenApp/). Coloring indicates transcripts per million (TPM) per tissue normalized to the average expression as indicated in the legend. The size of the dots represents the percentage of cells of this cell type expressing the gene. Coelomocyte data is highlighted by the blue box. (B) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb. MCA-3b is expressed in GABAergic neurons of flwr-1(ok3128) mutant animals (unc-47 promoter). Two-way ANOVA with Tukey’s correction. N = 3 biological replicates. ns not significant p > 0.05, *** p < 0.001.

FLWR-1 structure prediction does not imply Ca2+ conducting ability but PI(4,5)P2 binding

Earlier version: We observed increased rather than decreased Ca2+ influx during stimulation, which contrasts findings from D. melanogaster Flower (Yao et al., 2017). Previously, an evolutionarily conserved glutamate residue (E78 in the Drosophila protein) within the transmembrane domain, which was found to be essential for Drosophila Flower function, was suggested to represent a Ca2+ selectivity filter because of similarities to Cav1.2 and TRP channels (Yao et al., 2009). This group proposed that Flower, like TRP channels, can form homo-tetramers to conduct Ca2+ (Hoenderop et al., 2003; Yao et al., 2009; Zhang et al., 2023). We therefore wondered whether loss of the conserved glutamate residue impacts FLWR-1 function in C. elegans (Fig. 8A). Expressing a mutant variant of FLWR-1 in which glutamate 74 is exchanged to glutamine (E74Q) significantly decreased the aldicarb resistance of the flwr-1 mutant (Fig. 8B), yet did not fully rescue it. This suggests that the respective glutamate is not essential for C. elegans FLWR-1 function. To get more insight into the putative structure and oligomerization of FLWR-1, we used Alphafold 3 (AF3) predictions to estimate the approximate location of the conserved glutamate residue within a putative FLWR-1 homo-tetramer (Abramson et al., 2024; Evans et al., 2022). One of the predicted complexes indeed revealed a pore-like structure (Figure 8 – Figure supplement 1A). However, while conventional Ca2+-channels contain a glutamate residue within their pore domain (Chen et al., 2023), E74 of FLWR-1 is not predicted to be part of pore-lining residues in the putative tetramer (Figure 8 – Figure supplement 1B).

Revised version: We observed increased rather than decreased Ca2+ levels during stimulation, which contrasts findings from D. melanogaster Flower (Yao et al., 2017). Previously, an evolutionarily conserved glutamate residue (E78 in the Drosophila protein) within the transmembrane domain, which was found to be essential for Drosophila Flower function, was suggested to represent a Ca2+ selectivity filter because of similarities to Cav1.2 and TRP channels (Yao et al., 2009). This group proposed that Flower, like TRP channels, can form homo-tetramers to conduct Ca2+ (Hoenderop et al., 2003; Yao et al., 2009; Zhang et al., 2023). We therefore wondered whether loss of the conserved glutamate residue impacts FLWR-1 function in C. elegans (Fig. 8A). Expressing a mutant variant of FLWR-1 in which glutamate 74 is exchanged to glutamine (E74Q) significantly decreased the aldicarb resistance of the flwr-1 mutant (Fig. 8B), yet did not fully rescue it. This suggests that the respective glutamate is not essential for C. elegans FLWR-1 function. To get more insight into the putative structure and oligomerization of FLWR-1, we used Alphafold 3 (AF3) predictions to estimate the approximate location of the conserved glutamate residue within a putative FLWR-1 homo-tetramer (Abramson et al., 2024; Evans et al., 2022). One of the predicted complexes indeed revealed a pore-like structure (Figure 8 – Figure supplement 1A). However, while conventional Ca2+-channels contain a glutamate residue within their pore domain (Chen et al., 2023), E74 of FLWR-1 is not predicted to be part of pore-lining residues in the putative tetramer (Figure 8 – Figure supplement 1B).

Basic amino acid residues on the intracellular surface of FLWR-1 may be involved in PI(4,5)P2 lipid binding

Earlier version: (A) Partial alignment of the amino acid sequences of FLWR-1, hFwe4 (Homo sapiens), and Fwe-Ubi (Drosophila melanogaster). Shading depicts evolutionary conservation of amino acid residues (black – identity; grey - homology). (B) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb. Two-way ANOVA with Tukey’s correction. Selected comparisons are depicted. N = 3 biological replicates. *** p < 0.001. (C) Exemplary images of CCs in wild type and flwr-1(ok3128) animals expressing GFP fused to the PH-domain of PLCδin CCs. Scale bar, 10 µm. (D) Mean (± SEM) corrected GFP fluorescence. Membrane fluorescence was estimated by measuring the perimeter of CCs and subtracting the fluorescence in the interior of cells. Each dot indicates a single CC. a.u. = arbitrary units of fluorescence intensity. Kruskal-Wallis test. Only statistically significant differences are depicted. * p < 0.05, ** p < 0.01, *** p < 0.001. Number of animals imaged in N = 3 biological replicates: wild type = 27, flwr-1 = 25.

Revised version: (A) Partial alignment of the amino acid sequences of FLWR-1, hFwe4 (Homo sapiens), and Fwe-Ubi/FweA (Drosophila melanogaster). Shading depicts evolutionary conservation of amino acid residues (black – identity; grey - homology). (B) Mean (± SEM) fraction of moving animals after exposure to 1.5 mM aldicarb. Two-way ANOVA with Tukey’s correction. Selected comparisons are depicted. N = 3 biological replicates. *** p < 0.001. (C) Exemplary images of CCs in wild type and flwr-1(ok3128) animals expressing GFP fused to the PH-domain of PLCδin CCs. Scale bar, 10 µm. (D) Mean (± SEM) corrected GFP fluorescence. Membrane fluorescence was estimated by measuring the perimeter of CCs and subtracting the fluorescence in the interior of cells. Each dot indicates a single CC. a.u. = arbitrary units of fluorescence intensity. Kruskal-Wallis test. Only statistically significant differences are depicted. * p < 0.05, ** p < 0.01, *** p < 0.001. Number of animals imaged in N = 3 biological replicates: wild type = 27, flwr-1 = 25.

Removed in the revised version: Structural analysis of important amino acid residues in FLWR-1 and other proteins.

Removed in the revised version: (A) Top view of an Alphafold3 prediction of a tetrameric FLWR-1 structure. Glutamate 74 is indicated as colored spheres in each monomer. (B) Cryo-EM structure of the human L-type voltage-gated calcium channel Cav1.2 (PDB: 8EOG) (Chen et al., 2023). Glutamate residues in the central pore are indicated as spheres. Ca2+ ions within the pore are depicted as turquoise spheres. (C) X-Ray structure of the Kir2.2 potassium channel in a complex with PI(4,5)P2 (PDB: 3SPI) (Hansen et al., 2011). Basic amino acid residues in the proximity of PI(4,5)P2 are depicted as sticks. (D) PI(4,5)P2 was docked to the Alphafold3 structure of FLWR-1 using AutoDock Vina (Eberhardt et al., 2021). Basic amino acid residues in the proximity of PI(4,5)P2 are depicted as sticks. Arginine 27 and lysine 31 are indicated.

Removed in the revised version: FLWR-1 is physically close to MCA-3 in the plasma membrane of BWMs and promotes synaptic localization of MCA-3.

Removed in the revised version: (A) Schematic representation of the BiFC assay. (B) Representative micrographs of animals expressing FLWR-1 fused to one fragment of mVenus (VN173) and either MCA-3b, NRA-2/Nicalin or UNC-1 stomatin, interacting with gap junctions, respectively, fused to the other fragment of mVenus (VC155). Scale bar, 10 µm. (C) Representative micrographs of the DNC of animals expressing MCA-3b::YFP in cholinergic neurons (unc-17 promoter). Scale bar, 5 µm. (D, E) Analysis of MCA-3b::YFP fluorescence (mean ± SEM) along the nerve cord (D) and as the ratio of fluorescence in synaptic puncta over fluorescence in the inter-puncta regions (E). Each dot represents a single animal. Number of animals imaged in N = 3 biological replicates: wild type = 35, flwr-1 = 35. Unpaired t-test. ** p < 0.001.

Removed in the revised version: Expression of FLWR-1 and MCA-3b for BiFC does not affect locomotion behavior:

Removed in the revised version: Mean (± SEM) swimming cycles of animals analyzed by BiFC as depicted in Fig. 8B. N2 wild type animals were used as a control strain, as well as flwr-1(ok3128) and mca-3(ok2048) mutants. Number of animals accumulated from N = 4 (wild type) or N = 2 (all other strains) biological replicates: wild type = 90, FLWR-1/MCA-3b = 16, FLWR- 1/NRA-2 = 29, flwr-1 = 40, mca-3 = 42. One-way ANOVA with Tukey’s correction. ns not significant p > 0.05, *** p < 0.001.

Flower, like MCA-3, was also suggested to bind to phosphatidylinositol-4,5-bisphosphate (PI(4,5)P2), an important regulator of endo/exocytosis and thus SV recycling, and to affect its levels within the PM (Bednarek et al., 2007; Li et al., 2020; Lopreiato et al., 2014). Binding of transmembrane proteins (including TRP channels) to PI(4,5)P2 is generally mediated by positively charged amino acid residues (lysine and arginine) close to the intracellular side of the PM, which bind to the negatively charged phosphate residues (Figure 8 – Figure supplement 1C; (Duncan et al., 2020; Hansen et al., 2011; Ribalet et al., 2005; Rohacs, 2024). The predicted FLWR-1 structure similarly contains basic amino acids which are likely to be close to the intracellular PM leaflet (Figure 8 – Figure supplement 1D). We used in silico docking prediction to estimate how PI(4,5)P2 may bind to FLWR-1. In the resulting structure, both phosphate residues are close to lysine 31 and arginine 27 (Figure 8 – Figure supplement 1D). A positive charge at this position is evolutionarily conserved, suggesting an important function (Fig. 8A). Moreover, these amino acids have been suggested to mediate PI(4,5)P2 binding in D. melanogaster Flower (Li et al., 2020). Expression of a mutated version of FLWR-1 in which both residues were replaced by alanine only partially rescued aldicarb resistance (Fig. 8B), implying that FLWR-1 function may be facilitated by these residues and their potential interaction with PI(4,5)P2. These findings implicate that the conserved E74 is not essential, yet still important for FLWR-1 activity. In addition, an interaction with PI(4,5)P2 via basic amino acids also seems to be an important factor. However, we do note that FLWR-1 structural models as depicted here rely on in silico predictions. The actual biological structure of FLWR-1, or its homologues, has yet to be determined using X-ray crystallography or cryo-EM.

The pleckstrin homology (PH) domain of phospholipase Cδ (PLCδ) fused to GFP can be used to estimate PI(4,5)P2 levels within cells (Botelho et al., 2000; Chen et al., 2014; Stauffer et al., 1998). This was used to show that reduction of MCA-3 function induced decreased PH-PLCδ-GFP fluorescence levels in CCs (Bednarek et al., 2007). We thus used this assay to assess the role of FLWR-1 in endosomal PI(4,5)P2 levels in CCs. Mutation of flwr-1 reduced the observed GFP fluorescence suggesting that FLWR-1 positively affects PI(4,5)P2 levels (Fig. 8C, D). We further asked if the putative PI(4,5)P2 binding site mutations would affect binding of the reporter to the CC membrane. In fact, while expressing FLWR-1 specifically in CCs fully rescued the GFP levels at the CC membrane, expressing the FLWR-1(R27A/K31A) double mutant did not rescue the phenotype, while the E74Q mutant was not significantly different to wild type (Fig. 8C, D). These findings indicate that FLWR-1 function affects MCA-3 activity through an interaction via PI(4,5)P2 and that FLWR-1 may have an influence on PI(4,5)P2 levels.

A putative direct interaction of FLWR-1 and MCA-3 may increase MCA-3 function dynamically

Earlier version: Next, we asked how FLWR-1 may affect MCA-3 through PI(4,5)P2, and whether this may involve proximity of the two proteins. To test this, we used Bimolecular Fluorescence Complementation (BiFC) in vivo. This assay uses two non-fluorescent fragments of mVenus (VN173 and VC155, respectively), which can recombine and complement each other to form a fluorescent protein, as long as they come into close proximity within the cell (Almedom et al., 2009; Chen et al., 2007; Hu et al., 2002). The fragments can be fused to putative interaction partners and fluorescence complementation (including a covalent bond formation in the reconstituted fluorophore) verifies interaction. We thus fused the VC155 fragment to the C-terminus of FLWR-1 and the VN173 fragment to the C-terminus of MCA-3, and co-expressed both proteins (FLWR-1 was expressed from its own promoter, while MCA-3 was expressed in BWM cells; Fig. 9A). Fluorescence could readily be observed at the PM of muscles cells, but not in other tissues that also express FLWR-1 (Fig. 9B). As a control, we used two other proteins for probing interactions with FLWR-1: First, the ER-membrane protein NRA-2, which is a homologue of mammalian Nicalin, a component of the translocon-associated Nicalin-TMEM147-NOMO complex that assists folding and assembly of multipass transmembrane proteins in the ER (McGilvray et al., 2020; Smalinskaitė et al., 2022). We previously showed interactions of Nicalin/NRA-2 and the NOMO homologue NRA-4 by BiFC (Almedom et al., 2009). We would expect that FLWR-1, as a multipass TM protein, would also get into close proximity with the NRA-2/Nicalin protein upon biogenesis. Indeed, a clear ER-localized signal was observed when probing FLWR-1::VC155 and NRA-2::VN173 interactions in muscle (Fig. 9. B). Second, to exclude that FLWR-1 would interact with any protein presented in the BiFC assay, we used the UNC-1 stomatin, that is associated with innexin gap junctions in neurons and muscles (Chen et al., 2007). Co-expressing FLWR-1::VC155 with UNC-1::VN173 in BWMs, however, did not lead to distinct GFP reconstitution, apart from some general dim fluorescence that was hardly above the level of autofluorescence (Fig. 9B). These findings indicate that FLWR-1 and MCA-3 may interact with each other in the PM. If this interaction is functional, linking the two proteins via BiFC may not have adverse effects on either protein. We thus tested the transgenic animals in swimming assays (Figure 9 – Figure supplement 1). Neither expression of FLWR-1 with MCA-3 or FLWR-1 with NRA-2 affected swimming locomotion, showing that there were no dominant negative effects. Note that wild type FLWR-1 and MCA-3 were present in these animals. However, as mca-3 mutants have a significant reduction in swimming ability (Figure 9 – Figure supplement 1), the fact that FLWR-1/MCA-3 BiFC had no adverse effects, indicates that the physical interaction of the two proteins is in line with their native function.

Revised version: Next, we asked how FLWR-1 may affect MCA-3 through PI(4,5)P2, and whether this may involve proximity of the two proteins. To test this, we used Bimolecular Fluorescence Complementation (BiFC) in vivo. This assay uses two non-fluorescent fragments of mVenus (VN173 and VC155, respectively), which can recombine and complement each other to form a fluorescent protein, as long as they come into close proximity within the cell (Almedom et al., 2009; Chen et al., 2007; Hu et al., 2002). The fragments can be fused to putative interaction partners and fluorescence complementation (including a covalent bond formation in the reconstituted fluorophore) verifies interaction. We thus fused the VC155 fragment to the C-terminus of FLWR-1 and the VN173 fragment to the C-terminus of MCA-3, and co-expressed both proteins (FLWR-1 was expressed from its own promoter, while MCA-3 was expressed in BWM cells; Fig. 9A). Fluorescence could readily be observed at the PM of muscles cells, but not in other tissues that also express FLWR-1 (Fig. 9B). As a control, we used two other proteins for probing interactions with FLWR-1: First, the ER-membrane protein NRA-2, which is a homologue of mammalian Nicalin, a component of the translocon-associated Nicalin-TMEM147-NOMO complex that assists folding and assembly of multipass transmembrane proteins in the ER (McGilvray et al., 2020; Smalinskaitė et al., 2022). We previously showed interactions of Nicalin/NRA-2 and the NOMO homologue NRA-4 by BiFC (Almedom et al., 2009). We would expect that FLWR-1, as a multipass TM protein, would also get into close proximity with the NRA-2/Nicalin protein upon biogenesis. Indeed, a clear ER-localized signal was observed when probing FLWR-1::VC155 and NRA-2::VN173 interactions in muscle (Fig. 9. B). Second, to exclude that FLWR-1 would interact with any protein presented in the BiFC assay, we used the UNC-1 stomatin, that is associated with innexin gap junctions in neurons and muscles (Chen et al., 2007). Co-expressing FLWR-1::VC155 with UNC-1::VN173 in BWMs, however, did not lead to distinct YFP reconstitution, apart from some general dim fluorescence that was hardly above the level of autofluorescence (Fig. 9B). These findings indicate that FLWR-1 and MCA-3 may interact with each other in the PM. If this interaction is functional, linking the two proteins via BiFC may not have adverse effects on either protein. We thus tested the transgenic animals in swimming assays (Figure 9 – Figure supplement 1). Neither expression of FLWR-1 with MCA-3 or FLWR-1 with NRA-2 affected swimming locomotion, showing that there were no dominant negative effects. Note that wild type FLWR-1 and MCA-3 were present in these animals. However, as mca-3 mutants have a significant reduction in swimming ability (Figure 9 – Figure supplement 1), the fact that FLWR-1/MCA-3 BiFC had no adverse effects, indicates that the physical interaction of the two proteins is in line with their native function.

Added in the revised version: FLWR-1 is physically close to MCA-3 in the plasma membrane of BWMs and promotes synaptic localization of MCA-3

Added in the revised version: (A) Schematic representation of the BiFC assay. (B) Representative micrographs of animals expressing FLWR-1 fused to one fragment of mVenus (VN173) and either MCA-3b, NRA-2/Nicalin or UNC-1 stomatin, interacting with gap junctions, respectively, fused to the other fragment of mVenus (VC155). Scale bar, 10 µm. (C) Representative micrographs of the DNC of animals expressing MCA-3b::YFP in cholinergic neurons (unc-17 promoter). Scale bar, 5 µm. (D, E) Analysis of MCA-3b::YFP fluorescence (mean ± SEM) along the nerve cord (D) and as the ratio of fluorescence in synaptic puncta over fluorescence in the inter-puncta regions (E). Each dot represents a single animal. Number of animals imaged in N = 3 biological replicates: wild type = 35, flwr-1 = 35. Unpaired t-test. ** p < 0.001.

These observations indicate that MCA-3 and FLWR-1 may be interacting directly, and thus their function may jointly be responsible for the Ca2+-induced recycling of SVs. Since the increased Ca2+ levels in neurons upon stimulation in flwr-1 mutants may reflect altered MCA-3 abundance at synapses, we tested if FLWR-1 would affect localization of MCA-3 along neuronal membranes. We expressed MCA-3b::YFP in cholinergic neurons, which was observable along axonal membranes and was enriched in synaptic specializations (Fig. 9C). In flwr-1 mutants, this synaptic fluorescence was not overall altered (Fig. 9D), however, when we compared the YFP signals within the synaptic specializations relative to the inter-synaptic axon shafts, this demonstrated a significantly reduced presence of MCA-3 in synaptic puncta and a relocation to the axonal membranes (Fig. 9E). This is in line with the idea that FLWR-1 function may stabilize / augment MCA-3 expression in the plasma membrane and / or its localization in synaptic puncta, to promote efficient synaptic function and SV recycling.

Discussion

Earlier version: Drosophila Flower was shown to affect SV recycling and to alter synaptic Ca2+ levels upon stimulation. It was suggested to form a Ca2+ channel that gets inserted into the plasma membrane upon SV fusion. Here, we showed that also in C. elegans, FLWR-1 affects synaptic Ca2+ levels following stimulation, that FLWR-1 is localized to SVs in neuronal cells, and that in its absence, endosomal structures with slowed acidification kinetics accumulate, suggesting defective recovery of SVs from recycling endosomes. FLWR-1 may affect SV recycling by stimulating the PMCA MCA-3, possibly by directly interacting with it, and via modulation of PI(4,5)P2 levels, may affect the abundance of MCA-3 in the synaptic membrane. We found that loss of FLWR-1 conveyed increased Ca2+ influx in motor neurons, particularly GABAergic neurons, thus leading to more GABA release, an E/I imbalance at the NMJ, and leading to aldicarb resistance. The function of FLWR-1 appears to be ubiquitously involved in membrane trafficking and PM-endosomal recycling, as we found it to affect endocytosis and PM PI(4,5)P2 levels also in coelomocytes. Our findings are summarized in a model in Fig. 10.

Revised version: Drosophila Flower was shown to affect SV recycling and to alter synaptic Ca2+ levels upon stimulation. It was suggested to form a Ca2+ channel that gets inserted into the plasma membrane upon SV fusion. Here, we showed that also in C. elegans, FLWR-1 affects synaptic Ca2+ levels following stimulation, that FLWR-1 is localized to SVs in neuronal cells, and that in its absence, endosomal structures with slowed acidification kinetics accumulate, suggesting defective recovery of SVs from recycling endosomes. FLWR-1 may affect SV recycling by stimulating the PMCA MCA-3, possibly by directly interacting with it, and via modulation of PI(4,5)P2 levels, may affect the abundance of MCA-3 in the synaptic membrane. We found that loss of FLWR-1 conveyed increased Ca2+ entry into the motor neuron cytosol, particularly GABAergic neurons, thus leading to more GABA release, an E/I imbalance at the NMJ, and leading to aldicarb resistance. Since the C. elegans NMJ comprises cholinergic and GABAergic neurons, and the aldicarb assay relies on muscle contraction, interpreting results is complicated if a given mutation affects the two motor neuron types (or muscle physiology) differently. We thus included cell-type specific rescue in our analyses. The function of FLWR-1 appears to be ubiquitously involved in membrane trafficking and PM-endosomal recycling, as we found it to affect endocytosis and PM PI(4,5)P2 levels also in coelomocytes. Our findings are summarized in a model in Fig. 10.

Earlier version: Model summarizing findings made in this study. For details, see discussion.

Revised version: Model summarizing findings made in this study

Added in the revised version: For details, see discussion.

Earlier version: MCA-3 was shown to facilitate endocytosis in CCs (Bednarek et al., 2007), indicating that both, increased Ca2+ levels as well as defective endocytosis may be influenced by deregulation / lack of stimulation of MCA-3 in flwr-1 mutants. Whether FLWR-1 directly modulates localization or activates function of MCA-3, or whether this is caused by a homeostatic change in the absence of FLWR-1, is unclear as yet. Our finding of a putative direct interaction of FLWR-1 with MCA-3 suggests that FLWR-1 could be inserted into the PM to stimulate activity of MCA-3 already present in the PM. Alternatively, MCA-3 may undergo dynamic PM insertion and endocytic recycling as part of SVs, along with FLWR-1 (mammalian PMCA was indeed found in purified SVs; (Takamori et al., 2006). In synapses, this would automatically happen upon SV fusion, such that local Ca2+ extrusion at release sites would ensure synapse functionality during intense activity. Such a function of PMCA was shown in Drosophila, where PMCA domains delineate areas of SV fusion from areas of SV endocytosis (Krick et al., 2021b). Thus, independent Ca2+ signaling in these two domains, via the P/Q type Ca2+ channel Cav2 (UNC-2 in C. elegans) and the L-type Cav1 (EGL-19 in C. elegans), mediating fusion and endocytosis, respectively, are facilitated by the PMCA. MCA-3 could be acutely regulated by FLWR-1, in addition to its regulation by phospholipids (Lopreiato et al., 2014).

Revised version: MCA-3 was shown to facilitate endocytosis in CCs (Bednarek et al., 2007), indicating that both, increased Ca2+ levels as well as defective endocytosis may be influenced by deregulation / lack of stimulation of MCA-3 in flwr-1 mutants. Whether FLWR-1 directly modulates localization or activates function of MCA-3, or whether this is caused by a homeostatic change in the absence of FLWR-1, is unclear as yet. Our finding of a putative direct interaction of FLWR-1 with MCA-3 suggests that FLWR-1 could be inserted into the PM to stimulate activity of MCA-3 already present in the PM. Alternatively, MCA-3 may undergo dynamic PM insertion and endocytic recycling as part of SVs, along with FLWR-1 (mammalian PMCA was indeed found in purified SVs; (Takamori et al., 2006). In synapses, this would automatically happen upon SV fusion, such that local Ca2+ extrusion at release sites would ensure synapse functionality during intense activity. Such a function of PMCA was shown in Drosophila, where PMCA domains delineate areas of SV fusion from areas of SV endocytosis (Krick et al., 2021). Thus, independent Ca2+ signaling in these two domains, via the P/Q type Ca2+ channel Cav2 (UNC-2 in C. elegans) and the L-type Cav1 (EGL-19 in C. elegans), mediating fusion and endocytosis, respectively, are facilitated by the PMCA. MCA-3 could be acutely regulated by FLWR-1, in addition to its regulation by phospholipids (Lopreiato et al., 2014). Proof of a direct physical interaction will have to be confirmed by biochemical assays and a biological structure of the putative FLWR-1/MCA-3 complex.

Removed in the revised version: Proof of a direct physical interaction will have to be confirmed by biochemical assays and a biological structure of the putative FLWR-1/MCA-3 complex.

Drosophila Flower was shown to increase neuronal PI(4,5)P2 levels, driving ADBE during sustained transmission (Li et al., 2020). We could show that an involvement with PI(4,5)P2 is likely conserved in nematodes, requiring the putative PI(4,5)P2 binding site residues R27 and K31 of FLWR-1. Similarly, MCA-3 was important to maintain PI(4,5)P2 in CCs (Bednarek et al., 2007). Thus, FLWR-1 and MCA-3 may have a shared role in regulating PI(4,5)P2 levels, or FLWR-1 PM insertion stabilizes local PI(4,5)P2. Recently, another SV protein, synaptotagmin1 (Syt1), it was shown to influence synaptic PI(4,5)P2 levels following inserts fusion, by recruiting PIP kinase I gamma (PIPKIγ) (Bolz et al., 2023). Syt1 provides a binding site for PIPKIγ, involving R322 and K326, as part of the sequence RLKKK. In FLWR-1, R27 and K31 are part of the sequence RFLAK.

Earlier version: Our results indicate that MCA-3 may be functionally affected / not stimulated in flwr-1 mutants, explaining increased Ca2+ levels during neuronal stimulation (Brini and Carafoli, 2011; Chamberland et al., 2019; Ono et al., 2019). The increased Ca2+ influx we observed was at first surprising, since work in Drosophila yielded opposite results during 40 Hz electrical stimulation in neuromuscular boutons of Flower mutants (Yao et al., 2017). Upon stimulation with 10 Hz, however, these animals showed an increased GCaMP signal compared to wild type, in line with our observations. Potentially, increased Ca2+ in flwr-1 mutants is only observed with lower, more physiological stimulation. Stimulation by ChrimsonSA, which we used in the Ca2+ imaging assays, is not as vigorous, thus we are likely in a regime comparable to the 10 Hz stimulation in Drosophila. Using stronger stimulation via ChR2(H134R), as in the electrophysiological assays, led to a prominent rundown of evoked amplitudes. This may be a consequence of lower Ca2+ levels, and interference of MCA-3 with UNC- 2 dependent Ca2+ microdomains required for exocytosis, leading to less fusion events. During intense stimulation, larger amounts of PMCA/MCA-3 may be inserted into the PM due to the high number of SVs fusing. Without FLWR-1, less PI(4,5)P2 is generated and thus MCA-3 may not be effectively recycled from the PM, thus leading to reduced Ca2+ levels as observed in Drosophila.

Revised version: Our results indicate that MCA-3 may be functionally affected / not stimulated in flwr-1 mutants, explaining increased Ca2+ levels during neuronal stimulation (Brini and Carafoli, 2011; Chamberland et al., 2019; Ono et al., 2019). The increased Ca2+ level increase we observed was at first surprising, since work in Drosophila yielded opposite results during 40 Hz electrical stimulation in neuromuscular boutons of Flower mutants (Yao et al., 2017). Upon stimulation with 10 Hz, however, these animals showed an increased GCaMP signal compared to wild type, in line with our observations. Potentially, increased Ca2+ in flwr-1 mutants is only observed with lower, more physiological stimulation. Stimulation by ChrimsonSA, which we used in the Ca2+ imaging assays, is not as vigorous, thus we are likely in a regime comparable to the 10 Hz stimulation in Drosophila. Using stronger stimulation via ChR2(H134R), as in the electrophysiological assays, led to a prominent rundown of evoked amplitudes. This may be a consequence of lower Ca2+ levels, and interference of MCA-3 with UNC-2 dependent Ca2+ microdomains required for exocytosis, leading to less fusion events. During intense stimulation, larger amounts of PMCA/MCA-3 may be inserted into the PM due to the high number of SVs fusing. Without FLWR-1, less PI(4,5)P2 is generated and thus MCA-3 may not be effectively recycled from the PM, thus leading to reduced Ca2+ levels as observed in Drosophila.

Earlier version: Even though we observed increased instead of decreased Ca2+ influx in flwr-1 mutants, in line with a stimulatory effect of FLWR-1 on MCA-3, we cannot exclude a possible Ca2+-conductivity of FLWR-1. It was discussed that the kinetics of Ca2+ influx induced by Flower in Drosophila are too slow and the extent too little as to directly trigger known modes of endocytosis (Brose and Neher, 2009; Leitz and Kavalali, 2016; Xue et al., 2012), e.g. by activating the Ca2+ sensor calmodulin (Wu et al., 2009). To assess the possibility of FLWR-1 forming an ion channel we asked AlphaFold3 to model FLWR- 1 as a tetramer. The putative pore lining residues were mostly hydrophobic or apolar, thus not in agreement with a function as an ion channel. The suggested similarity to the selectivity filter of VGCCs, including the sequence EGW or EAW (in FLWR-1 this is EAP; Yao et al. (2009) was not confirmed in the AF3 models: While glutamate faces the extracellular end of the VGCC pore in all four channel modules, E74FLWR-1 pointed away from the pore. We thus think that FLWR-1 is unlikely to form an ion channel, and suggest that it facilitates SV recycling by stimulating MCA-3 activity.

Revised version: Even though we observed increased instead of decreased Ca2+ level increase in flwr-1 mutants, in line with a stimulatory effect of FLWR-1 on MCA-3, we cannot exclude a possible Ca2+-conductivity of FLWR-1. It was discussed that the kinetics of Ca2+ levels rising, induced by Flower in Drosophila, are too slow and the extent too little as to directly trigger known modes of endocytosis (Brose and Neher, 2009; Leitz and Kavalali, 2016; Xue et al., 2012), e.g. by activating the Ca2+ sensor calmodulin (Wu et al., 2009). To assess the possibility of FLWR-1 forming an ion channel we asked AlphaFold3 to model FLWR-1 as a tetramer. The putative pore lining residues were mostly hydrophobic or apolar, thus not in agreement with a function as an ion channel. The suggested similarity to the selectivity filter of VGCCs, including the sequence EGW or EAW (in FLWR-1 this is EAP; Yao et al., 2009) was not confirmed in the AF3 models: While glutamate faces the extracellular end of the VGCC pore in all four channel modules, E74FLWR-1 pointed away from the pore. We thus think that FLWR-1 is unlikely to form an ion channel, and suggest that it facilitates SV recycling by stimulating MCA-3 activity.

Earlier version: Activating PMCAs via PI(4,5)P2 (Berrocal et al., 2017) would remove Ca2+ and prevent sustained phospholipase C activity, thus protecting PI(4,5)P2 at the PM from hydrolysis (Lopreiato et al., 2014; Penniston et al., 2014), and activating different forms of endocytosis (Blumrich et al., 2023; Bolz et al., 2023; Posor et al., 2015). Such interactions could further trigger a positive feedback loop leading to invagination of the PM (Yao et al., 2017). Both PMCAs and Flower were suggested to directly bind PI(4,5)P2 using positively charged amino acid residues close to the membrane (Filoteo et al., 1992; Li et al., 2020). In line with this, the FLWR-1 structure model exhibits amino acids in suitable positions. This could lead to a close arrangement of both proteins in PI(4,5)P2 lipid microdomains (Katan and Cockcroft, 2020). Flower could also activate PMCAs in intracellular compartments. One PMCA variant was shown to mainly localize to recycling SVs and endosomes, and Ca2+ influx into these is an important factor in clearance from the pre-synapse (Ono et al., 2019). Overall, this may explain our finding of increased numbers of large endocytic structures in flwr-1 synapses. Effects of flwr-1 mutations on PI(4,5)P2 at the PM may negatively affect ADBE, while effects on PI(4,5)P2 in bulk endosomes may slow the breakdown of these structures (Fig. 10). Our ultrastructural analysis contrasts EM analyses in Drosophila which yielded fewer, rather than more, bulk endosomal structures in Flower mutants (Li et al., 2020; Yao et al., 2017). However, this was following prolonged (ten minute) stimulation, possibly resulting in different outcomes compared to our 30 s optogenetic stimulus. Alternatively, the subcellular localization of Flower function may differ between organisms.

Revised version: Activating PMCAs via PI(4,5)P2 (Berrocal et al., 2017) would remove Ca2+ and prevent sustained phospholipase C activity, thus protecting PI(4,5)P2 at the PM from hydrolysis (Lopreiato et al., 2014; Penniston et al., 2014), and activating different forms of endocytosis (Blumrich et al., 2023; Bolz et al., 2023; Posor et al., 2015). Such interactions could further trigger a positive feedback loop leading to invagination of the PM (Yao et al., 2017). Both PMCAs and Flower were suggested to directly bind PI(4,5)P2 using positively charged amino acid residues close to the membrane (Filoteo et al., 1992; Li et al., 2020). In line with this, the FLWR-1 structure model exhibits amino acids in suitable positions. This could lead to a close arrangement of both proteins in PI(4,5)P2 lipid microdomains (Katan and Cockcroft, 2020). Flower could also activate PMCAs in intracellular compartments. One PMCA variant was shown to mainly localize to recycling SVs and endosomes, and Ca2+ entry into these is an important factor in clearance from the pre-synapse (Ono et al., 2019). Overall, this may explain our finding of increased numbers of large endocytic structures in flwr-1 synapses. Effects of flwr-1 mutations on PI(4,5)P2 at the PM may negatively affect ADBE, while effects on PI(4,5)P2 in bulk endosomes may slow the breakdown of these structures (Fig. 10). Our ultrastructural analysis contrasts EM analyses in Drosophila which yielded fewer, rather than more, bulk endosomal structures in Flower mutants (Li et al., 2020; Yao et al., 2017). However, this was following prolonged (ten minute) stimulation, possibly resulting in different outcomes compared to our 30 s optogenetic stimulus. Alternatively, the subcellular localization of Flower function may differ between organisms.

Last, Flower could facilitate endocytosis through calcineurin (Yao et al., 2017), to stimulate bulk endosome formation. This Ca2+- and calmodulin-dependent phosphatase dephosphorylates endocytic proteins to accelerate SV recycling (Cousin and Robinson, 2001; Kumashiro et al., 2005; Yamashita, 2012). How could the increased Ca2+ levels in the absence of FLWR-1 and the decreased rate of pHluorin decay fit to this idea? Possibly, increased activation of calcineurin shifts the endocytic mode from ultrafast endocytosis (UFE) (Watanabe et al., 2013) to ADBE (Kittelmann et al., 2013b; Yamashita, 2012). Structures formed by bulk endocytosis are acidified more slowly than small endocytic vesicles generated by UFE, which could explain the delayed decay of pHluorin fluorescence (Gross and von Gersdorff, 2016; Okamoto et al., 2016). More pronounced neurotransmission, due to increased SV fusion, causes an increase in bulk invaginations (Clayton et al., 2008), yet the endocytic structures formed this way are not efficiently disassembled. The C. elegans calcineurin TAX-6 is further involved in activity-dependent bulk endocytic processes, yet in a different context, i.e. reorganization of GABAergic synapses during development (Cuentas-Condori et al., 2023; Miller-Fleming et al., 2016). The remote possibility that FLWR-1 also takes part in this process could imply that developmental aspects of FLWR-1 function may influence E/I balance in adult NMJs.

In sum, our work confirms a conserved role of FLWR-1 in neurotransmission and SV recycling even though we observed differences to data from Drosophila NMJs. We found an unexpected upregulation of neurotransmission in flwr-1 mutants which has not yet been observed in other animals, and which may result from FLWR-1 acting on PMCAs directly. The observed reduced Ca2+ levels in hyper-stimulated Drosophila synapses may instead result from ‘stranded’ PMCA in the PM, due to impaired recycling, thus increasing Ca2+ extrusion. Further investigations are required to confirm involvement of FLWR-1 in MCA-3 regulation and the nature of the relationship between these proteins.